Sld3CBD–Cdc45 structural insights into Cdc45 recruitment for CMG complex formation on DNA replication

Author response:

The following is the authors’ response to the previous reviews

Reviewer #1 (Public review):

Summary:

The crystal structure of the Sld3CBD-Cdc45 complex presented by Li et al. is a significant contribution that enhances our understanding of CMG formation during the rate-limiting step of DNA replication initiation. This structure provides crucial insights into the intermediate steps of CMG formation, and the particle analysis and model predictions compellingly describe the mechanism of Cdc45 loading. Building upon previously known Sld3 and Cdc45 structures, this study offers new perspectives on how Cdc45 is recruited to MCM DH through the Sld3-Sld7 complex. The most notable finding is the structural rearrangement of Sld3CBD upon Cdc45 binding, particularly the α8-helix conformation, which is essential for Cdc45 interaction and may also be relevant to its metazoan counterpart, Treslin. Additionally, the conformational shift in the DHHA1 domain of Cdc45 suggests a potential mechanism for its binding to Mcm2NTD. Furthermore, Sld3's ssDNA-binding experiments provide evidence of its novel functions in the DNA replication process in yeast, expanding our understanding of its role beyond Cdc45 recruitment.

Strengths:

The manuscript is generally well-written, with a precise structural analysis and a solid methodological section that will significantly advance future studies in the field. The predictions based on structural alignments are intriguing and provide a new direction for exploring CMG formation, potentially shaping the future of DNA replication research. This research also opens up several new opportunities to utilize structural biology to unravel the molecular details of the model presented in the paper.

Weaknesses:

The main weakness of the manuscript lies in the lack of detailed structural validation for the proposed Sld3-Sld7-Cdc45 model, and its CMG bound models, which could be done in the future using advanced structural biology techniques such as single particle cryo-electron microscopy. It would also be interesting to explore how Sld7 interacts with the MCM helicase, and this would help to build a detailed long-flexible model of Sld3-Sld7-Cdc45 binding to MCM DH and to show where Sld7 will lie on the structure. This will help us to understand how Sld7 functions in the complex. Also, future experiments would be needed to understand the molecular details of how Sld3 and Sld7 release from CMG is associated with ssARS1 binding.

The proposals based on this study provide new knowledge of the CMG formation process. We agree that our Sld3-Sld7-Cdc45 model will be further confirmed by cryo-EM. We improved our ssARS1-binding assay and quantified data (See the response to Recommendations for the authors of #3 review).

Reviewer #2 (Public review):

Summary

The manuscript presents valuable findings, particularly in the crystal structure of the Sld3CBD-Cdc45 interaction and the identification of additional sequences involved in their binding. The modeling of the Sld7-Sld3CBD-CDC45 subcomplex is novel, and the results provide insights into potential conformational changes that occur upon interaction. Although the single-stranded DNA binding data from Sld3 of different species is a minor weakness, the experiments support a model in which the release of Sld3 from the complex may be promoted by its binding to origin single-stranded DNA exposed by the helicase.

Strengths

The Sld3CBD-Cdc45 structure is a novel contribution, revealing critical residues involved in the interaction.

The model structures generated from the crystal data are well presented and provide valuable insights into the interaction sequences between Sld3 and Cdc45.

The experiments testing the requirements for interaction sequences are thorough and conducted well, with clear figures supporting the conclusions.

The conformational changes observed in Sld3 and Cdc45 upon binding are interesting and enhance our understanding of the interaction.

The modeling of the Sld7-Sld3CBD-CDC45 subcomplex is a new and valuable addition to the field.

The proposed model of Sld3 release from the complex through binding to single stranded DNA at the origin is intriguing.

Weaknesses

The section on the binding of Sld3 complexes to origin single-stranded DNA is somewhat weakened by the use of Sld3 proteins from different species. The comparisons between Sld3-CBD, Sld3CBD-Cdc45, and Sld7-Sld3CBD-Cdc45 involve complexes from different species, limiting the comparisons' value.

Although the study reveals that Sld3 binds to different residues of Cdc45 than those previously shown to bind Mcm or GINS, the data in the paper do not shed any additional light on how GINS and Sld3 binding to Cdc45 or Mcms. would affect each other. Other previous research has suggested that the binding of GINS and Sld3 to Mcm or Cdc45 may be mutually exclusive. The authors acknowledge that a structural investigation of Sld3, Sld7, Cdc45, and MCM during the stage of GINS recruitment will be a significant goal for future research.

We agree that it is better to use all samples from a source; however, due to limitations in protein expression, we used Sld7-Sld3CBD-Cdc45 from a different source. The two sources used in this study belong to the same family, and the proteins Sld7, Sld3 and Cdc45 share sequence conservation with similar structures predicted by Alphafold3 (RMSD = 0.356, 1.392, and 0.891 for Ca atoms of Sld7CTD, Sld7NTD-Sld3NTD, and Sld3CBD-Cdc45). Such similarity in source and proteins allows us to do the comparison. We also mentioned that a cryo-EM study of Sld3-Sld7-Cdc45-MCM and Sld3-Sld7-CMG structures will be a significant goal for future research in our manuscript.

Reviewer #3 (Public review):

Summary:

The paper by Li et al. describes the crystal structure of a complex of Sld3-Cdc45-binding domain (CBD) with Cdc45 and a model of the dimer of an Sld3-binding protein, Sld7, with two Sld3-CBD-Cdc45 for the tethering. In addition, the authors showed the genetic analysis of the amino acid substitution of residues of Sld3 in the interface with Cdc45 and biochemical analysis of the protein interaction between Sld3 and Cdc45 as well as DNA binding activity of Sld3 to the single-strand DNAs of the ARS sequence.

Strengths:

The authors provided a nice model of an intermediate step in the assembly of an active Cdc45-MCM-GINS (CMG) double hexamers at the replication origin, which is mediated by the Sld3-Sld7 complex. The dimer of the Sld3-Sld7 complexes tethers two MCM hexamers together for the recruitment of GINS-Pol epsilon on the replication origin.

Weaknesses:

The biochemical analysis should be carefully evaluated with more quantitative ways to strengthen the authors' conclusion even in the revised version.

In this revision, we improved our ssARS1-binding assay in more quantitative ways (See the response to Recommendations for the authors).

Reviewer #1 (Recommendations for the authors):

I thank the authors for all their replies to my previous questions and for doing all the necessary corrections. I am satisfied with most of their replies, however, upon second reading I have a few more suggestions which could help to improve the manuscript further and make an impact in the field. My comments are listed below.

(1) In general, the manuscript is well structured, but I feel that it requires professional English correction. In many places it was difficult to understand the sentences and I had to read it several times to understand it. Also, very long sentences should be avoided. The flow should be easy to read and understand, and that is why I feel it requires professional English correction.

Following the comment, we checked English carefully and shortened the very long sentences.

(2) Page 5, line 103, please include molecule after the word complex to make it like- "Only one complex molecule exists within an asymmetric unit."

We revised this sentence (P5/L103).

(3) Line 113- more than the N-terminal half of the protruding long helix α7 113 was disordered in the Sld3CBD-Cdc45 complex. This sentence is not clear. What does it mean more than the N-terminal half? Please rewrite it.

We revised this sentence to give the corresponding residue number “(D219–H231)” (P5/L114).

(4) Page 5, result 2- Conformation changes in Sld3CBD and Cdc45 for binding each other, this section may require a little restructuring. Line 130-131- "Therefore, the helix α8CTP seems to be an intrinsically disordered segment when Sld3 alone but 130 folds into a helix coupled to the binding partner Cdc45 in the Sld3CBD-Cdc45 complex." This statement is the crux of the structural finding and therefore, I feel it should move after the first sentence.

Thank you for your comments. We rewrote this part (P5/L128-131).

(5) Line 121-122: Compared to the isolated form (PDBIDs: 5DGO 121 for huCdc45 [31] and 6CC2 for EhCdc45 [33]) and the CMG form (PDBID: 3JC6. Write it in the same format. Make 6CC2 in bracket like other PDB IDs. Restructure this sentence.

We revised this sentence (P5/122-123).

(6) Line 127-129: This sentence is also not very clear.

We revised this sentence together with above No (4). (P5/L128-131)

(7) In my question 4- "Can authors add a supplementary figure showing the probability of disordernes..."., I meant to use a disorder prediction tool like IUPred for the protein sequences and show that α8 is predicted to be a disordered upon sequence analysis. This will help to show the inherent property of α8 helix, and it could add up to the understanding that a disordered region is being structured in the complex structure.

The structures showed that α8CTP is stabilized by binding with Cdc45, but disordered in Sld3CBD alone, indicating that this part is flexible, like an intrinsically disordered segment. We have deposited the structure to PDB, so predictions like IUPred cannot show meaningful information.

(8) Question 9 regarding Supplementary Figure 8- Please include your statement in the figure legend - "WT Sld3CBD was prepared in a complex with Cdc45, while the mutants of Sld3CBD existed alone, we calculated the elements of secondary structure from the crystal structure of Sld3CBD-Cdc45. The concentration of samples was controlled to the same level for CD measurement."

Following the comment, we optimized the figure legend of Supplementary Figure 8.

(9) Question 13- I understand that negative staining and SEC-SAXS experiments could be very tricky for such protein complexes, which have very long loops and are flexible. Did authors try a GraFix cross-linking before doing the negative staining TEM? If it is not being tried, then it might be a good idea to try it and it may help to get much cleaner particles and easier class averaging. Although I completely understand the technical challenges the authors describe and I agree with them, I still feel that one good experiment that shows this dimer model would be very helpful to strengthen the claim. I am concerned because if people start using a similar DLS experiment to calculate intermolecular distances, citing your paper, in many cases it might be a wrong interpretation. In case the negative staining still does not work, at least discuss your technical challenges in the discussion section and mention that SEC-SAXS showed a similar length of the complex and show the Guinier plot and Porod plots in the supplementary data.

We believe that DLS is one of the methods for analyzing the single particle size. Of course, the confirmation by multiple methods will give compelling evidence. Following the comment, we added SEC-SAXS data in the [Results] (P7/L194-196) (Cdc45 recruitment to MCM DH by Sld3 with partner Sld7) and Supplementary Figure 11. The Sld7-Sld3-Cdc45 forms a flexible, long shape. Each binding domain is rigid but linked by the long loops. The flexibility problems are caused by the long loop linkers, but not by binding. So, we did not try to use the cross-linking method for analysis experiments.

(10) Page 8, line 221- litter sequence specificity: Correct the word "litter" with little. Also, the word shaped is written as sharped at a few places in the manuscript. Please correct it.

We apologize for making such mistakes. We have modified these words.

(11) Page 9, line 237-238: Would it be possible to add a lane showing Sld7 binding to the ssDNA in figure 4. I recommend showing this to understand the ssDNA binding affinity of Sld7 by itself and it will also help us to compare when it is in complex with Sld3.

Considering that Sld7 on CMG is always a complex with Sld3, the ssDNA binding affinity should use the Sld3-Sld7 complex. Additionally, we attempted to overexpress Sld7, but could not obtain the target protein.

Reviewer #2 (Recommendations for the authors):

Thank you for the improved manuscript. The following sentence is unclear: "Cdc45 binds tighter to long ssDNA (>60 bases) with a litter sequence specificity".

We apologize for making such a mistake. We modified “litter” to “little”.

I found it challenging to understand which species were used while reading the results section and figure legends. I recommend that the authors revise the text in both the results and figure legends to clearly indicate when proteins from different species are being compared. Additionally, it would be valuable to explicitly acknowledge this limitation in the text.

Following the comment, we added a description for using different species in results (P8/L224-225) and figure legends (Supplementary Figure 14). We added more information in the Methods to explain why we used two species for preparing proteins.

Reviewer #3 (Recommendations for the authors):

Major points:

(1) The current title is not appropriate for the general readers. At least, DNA replication or DNA replication initiation should be added and abbreviations such as CBD should be avoided.

Following the comment, we added “DNA replication” into the title. Regarding “CBD”, since the full name of “Cdc45 binding domain” is too long, we continue to use Sld3CBD.

(2) As in my previous review, I asked for quantification of the EMSA assay shown in Figure 4 and Supplemental Figures 13 and 14. Since some signals of the bands are very weak, it is hard to conclude something. Given different protein concentrations used in the experiment, the authors should provide any kinds of value. For example, Sld3CBD-CDC45 shows weaker DNA binding than Sld3CBD alone (line 231). Is this true (or reproducible)? It is hard to conclude without any quantification.

We have repeated the EMSA assay four or more times with different rods of overexpression, purification and DNA synthesis, indicating that the EMSA assay is reproducible. In this revision, we changed the DNA stain and adjusted the ratio between the protein and ssDNA with increasing concentrations. The smeared bands of ssDNA with Sld7–Sld3ΔC–Cdc45 or Sld7–Sld3ΔC exhibit enhanced discernibility, and the ssDNA bands are intense enough for grayscale calculations (Figure 4 in the second revised version). We used a series of t-tests to confirm a significantly ssDNA residual level between Sld3CBD–Cdc45 to Sld3CBD, Sld7–Sld3ΔC–Cdc45, and Sld7–Sld3ΔCS (t-test, ****: P<0.0001). We also carefully controlled the sample amount in the EMAS assay and described it in the [Methods].

Moreover, in this EMSA assay (in Figure 4), the authors suggest that the disappearance of ssDNA bands corresponds with the binding of the protein to the DNA. However, it is also possible that the DNA is degraded. It is very important to show the band of protein-DNA complexes on the gel (a whole gel, not the parts of the gel shown in Figure). Why did the authors use this "insensitive" assay using SyberGreen, not radio-labelled ssDNA?

In this revision, we added a negative control of no ssDNA-binding by using ssARS1-3_3 for all protein samples (Sld3CBD, Sld3CBD–Cdc45, Sld7–Sld3ΔC–Cdc45 and Sld7–Sld3ΔC), which were the same rod of expression and purification for bound to ssARS1s (ssARS1-2 and ssARS1-5) (Figure 4), showing that the disappearance of ssDNA bands is caused by binding to proteins, not degradation. Moreover, this time, by changing the DNA stain and increasing the concentration of the samples, the smeared ssDNA bands exhibit enhanced discernibility in the high molecular weight regions when mixed with Sld7–Sld3ΔC–Cdc45 or Sld7–Sld3ΔC, whereas no bands appeared in the NC (ssARS1-3_1). The positions of smeared ssDNA bonds correspond to those of protein in the protein-stain pages, indicating that ssARS1 were complexed with proteins. Following the comment, we show all bands on the gel in Figure 4 and Supplementary Figure 14. Compared to Sld7–Sld3ΔC–Cdc45 or Sld7–Sld3ΔC, Sld3CBD and ssDNA bonds could not be observed because the pI value of Sld3CBD, which affects the entry of the samples into the gel.

We agree that using radio-labelled ssDNA can obtain a sensitive binding assay. However, current laboratory constraints did not allow us to use radio-labelled ssDNA. Furthermore, considering the characteristics of our target proteins, Sld3CBD, Sld3CBD–Cdc45, Sld7–Sld3ΔC–Cdc45, and Sld7–Sld3ΔC, we planned to perform the binding assay in a more natural state without any modifications, labelling or linkers. Additionally, we have attempted to use ITC experiments but failed in the measurements. Presumably, the conformational flexibility of Sld7-Sld3-Cdc45 and Sld7-Sld3 caused a thermodynamic anomaly.

Minor points:

(1) Line 215, 80b: This should be "80 nucleotides(nt)". Throughout the text, nucleotides is better than base to show the length of ssDNAs.

Thank you for your comments. We modified these words throughout the text.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: DNA replication requires Cdc45 and GINS recruitment into the MCM duplex hexamer by initiation factors to form an active helicase, the Cdc45–MCM–GINS (CMG) complex, at the replication origins. The initiation factor Sld3 is a central regulator of Cdc45 and GINS recruitment worked with Sld7 together. However, the mechanism through which Sld3 regulates CMG complex formation remains unclear. Here, we present the structure of the Sld3 Cdc45-binding-domain in complex with Cdc45 (Sld3CBD–Cdc45), showing detailed interactions and conformational changes required for binding to each other. The mutant analysis indicated that the binding between Sld3CBD and Cdc45 could be broken easily. We also revealed that Sld3CBD, GINS, and MCM bind to different sites on Cdc45 in the Sld3CDB–CMG model, indicating that after recruitment of Cdc45, Sld7–Sld3 could remain in Cdc45–MCM until CMG formation. The consistency between the particle size of Sld7–Sld3–Cdc45 and the distance between Sld3CBDs in the Cdc45–MCM dimer indicated the binding manner of the Cdc45–Sld3–[Sld7]2–Sld3–Cdc45 off/on MCM duplex hexamer. A DNA-binding assay of Sld3 and its complexes with single-stranded ARS1 fragments revealed a relationship between the dissociation of Sld7–Sld3 from CMG and the unwound single-stranded DNA. These findings help to further our understanding of the molecular basis of regulation of CMG complex formation by Sld3.

Revised version: DNA replication requires recruitment of Cdc45 and GINS into the MCM double hexamer by initiation factors to form an active helicase, the Cdc45–MCM–GINS (CMG) complex, at the replication origins. The initiation factor Sld3 is a central regulator of Cdc45 and GINS recruitment, working with Sld7 together. However, the mechanism through which Sld3 regulates CMG complex formation remains unclear. Here, we present the structure of the Sld3 Cdc45-binding domain in complex with Cdc45 (Sld3CBD–Cdc45), showing detailed interactions and conformational changes required for binding to each other. The mutant analysis indicated that the binding between Sld3CBD and Cdc45 could be broken easily. We also revealed that Sld3CBD, GINS, and MCM bind to different sites on Cdc45 in the Sld3CBD–CMG model, indicating that after recruitment of Cdc45, Sld7–Sld3 could remain in Cdc45–MCM until CMG formation. The consistency between the particle size of Sld7–Sld3–Cdc45 and the distance between Sld3CBDs in the Cdc45–MCM dimer indicated the binding manner of the Cdc45–Sld3–[Sld7]2–Sld3–Cdc45 off/on MCM double hexamer. A DNA-binding assay of Sld3 and its complexes with single-stranded ARS1 fragments revealed a relationship between the dissociation of Sld7–Sld3 from CMG and the unwound single-stranded DNA. These findings help to further our understanding of the molecular basis of the regulation of CMG complex formation by Sld3.

Introduction

Earlier version: Eukaryotic chromosomal DNA replication in the cell cycle, a tightly regulated process that ensures precise gene copying, begins with the unwinding of double-stranded DNA (dsDNA) into signal-stranded DNA (ssDNA), including the formation of an active helicase Cdc45–MCM–GINS (CMG) complex [1,2]. In the G1 phase of the yeast cell cycle, two inactive helicase Mcm2–7 hexamer rings (Mcm2–6–4–7–3–5) are loaded onto autonomously replicating sequences (ARSs) of dsDNA to form a double hexamer (MCM DH) at the replication origin [3–5]. To facilitate this loading, Cdt1 binds to the Mcm2–7 hexamer ring and stabilizes a gap formation between the Mcm2 and Mcm5, which serves as a dsDNA entry gate [6–9]. Next, the MCM DH on dsDNA is phosphorylated by Dbf4-dependent protein kinase (Cdc7 kinase or DDK) (Supplementary Figure 1A), allowing the key factor Sld3 and its partner Sld7 to recruit Cdc45 (Sld7–Sld3–Cdc45) to each MCM (Supplementary Figure 1B) [10–13]. Subsequently, the phosphorylation of Sld3 by cyclin-dependent kinase (CDK) regulates the assembly of GINS onto MCM DH with Dpb11, CDK-phosphorylated Sld2, and DNA polymerase ε to form CMG (Supplementary Figure 1C) [14–19]. Finally, the factors Sld2, Sld3, Sld7, and Dpb11 are released by elusive mechanisms from an active CMG and bidirectional replication by translocating along the 3′ to 5′ direction of the DNA strand (Supplementary Figure 1D) is facilitated [20, 21].

Revised version: Eukaryotic chromosomal DNA replication in the cell cycle, a tightly regulated process that ensures precise gene copying, begins with the unwinding of double-stranded DNA (dsDNA) into single-stranded DNA (ssDNA), including the formation of an active helicase Cdc45–MCM–GINS (CMG) complex [1,2]. In the G1 phase of the yeast cell cycle, two inactive helicase Mcm2–7 hexamer rings (Mcm2–6–4–7–3–5) are loaded onto autonomously replicating sequences (ARSs) of dsDNA to form a double hexamer (MCM DH) at the replication origin [3–5]. To facilitate this loading, Cdt1 binds to the Mcm2–7 hexamer ring and stabilizes a gap formation between the Mcm2 and Mcm5, which serves as a dsDNA entry gate [6–9]. Next, the MCM DH on dsDNA is phosphorylated by Dbf4-dependent protein kinase (DDK, also known as Cdc7 kinase) (Supplementary Figure 1A), allowing the key factor Sld3 and its partner Sld7 to recruit Cdc45 (Sld7–Sld3–Cdc45) to each MCM (Supplementary Figure 1B) [10–13]. Subsequently, the phosphorylation of Sld3 by cyclin-dependent kinase (CDK) regulates GINS assembly onto MCM DH with Dpb11, CDK-phosphorylated Sld2, and DNA polymerase ε to form CMG (Supplementary Figure 1C) [14–19]. Finally, the factors Sld2, Sld3, Sld7, and Dpb11 are released by elusive mechanisms from an active CMG, and bidirectional replication by translocating along the 3′ to 5′ direction of the DNA strand (Supplementary Figure 1D) is facilitated [20, 21].

Earlier version: As a central regulator of CMG formation in Saccharomyces cerevisiae, Sld3 functions as a bridge protein in the Sld7–Sld3–Cdc45 complex to recruit Cdc45 to DDK-phosphorylated MCM DH. This Sld3 contains three domains each of which binds mainly to one of Sld7, Cdc45, and MCM. The N-terminal domain of Sld3 (Sld3NTD: M1–L116) binds to the N-terminal domain (M1–D130) of Sld7 (Sld7NTD) (Sld7 C-terminal domain [Sld7CTD: K168–S257] is a self-dimerization domain) [22]. Following the NTD of Sld3, the middle part is a central portion known as the Cdc45-binding domain (Sld3CBD: S148–K430) [12, 23]. A previous study also demonstrated that a region (L510–R530) in the Sld3 C-terminal domain (Sld3CTD: T445–T679) binds to Mcm4 and 6 [24]. CDK-dependent phosphorylation of two residues (T600 and S622) downstream of Sld3CTD initiates the binding of GINS-carrying Dpb11 to Cdc45–MCM [19, 25]. In addition to its role as a bridge in the recruitment of Cdc45 and GINS, Sld3 binds to two single-stranded DNA (ssDNA) fragments of ARS1 which was identified as an origin of DNA replication (ssARS1-2 and ssARS1-5), but not to the corresponding double-stranded ARS1 (dsDNA) [26]. The specific association between Sld3 and ssDNA is not affected by CDK phosphorylation. Furthermore, the structures of the Sld7CTD dimer (PDBID:3X38) [22], Sld7NTD–Sld3NTD (PDBID:3X37) [22], Sld3CBD (PDBID:3WI3) [23], MCM DH (6F0L) [27], CMG (PDBIDs:3JC5, 3JC6, 3JC7) [28], CMG-DNA-polLJ (PDBID:7Z13) [29], CMG-DONSON-DNA (PDBID:8Q6O) [30], and so on, were determined through crystallography and cryogenic electron microscopy (cryo-EM). Cdc45 belongs to the DHH superfamily of proteins defined by the conserved triad motif DHH (Asp-His-His), and there is a DHH-associate domain (DHHA1: R523–L650) at the C-terminus [31]. Recent single-molecule biochemical assays have reported the stepwise recruitment of multiple Cdc45s to the MCM DH [32]. However, how Sld3–Sld7 brings Cdc45 into the MCM for the formation of CMG to regulate the initiation of DNA replication remains unclear.

Revised version: As a central regulator of CMG formation in Saccharomyces cerevisiae, Sld3 functions as a bridge protein in the Sld7–Sld3–Cdc45 complex to recruit Cdc45 to DDK-phosphorylated MCM DH. This Sld3 contains three domains, each of which binds mainly to one of Sld7, Cdc45, and MCM. The N-terminal domain of Sld3 (Sld3NTD: M1–L116) binds to the N-terminal domain (M1– D130) of Sld7 (Sld7NTD) (Sld7 C-terminal domain [Sld7CTD: K168–S257] is a self-dimerization domain) [22]. Following the NTD of Sld3, the middle part is a central portion, the Cdc45-binding domain (Sld3CBD: S148–K430) [12, 23]. A previous study also demonstrated that a region (L510–R530) in the Sld3 C-terminal domain (Sld3CTD: T445–T679) binds to Mcm4 and 6 [24]. CDK-dependent phosphorylation of two residues (T600 and S622) downstream of Sld3CTD initiates the binding of GINS-carrying Dpb11 to Cdc45–MCM [19, 25]. In addition to its role as a bridge in the recruitment of Cdc45 and GINS, Sld3 binds to two single-stranded DNA (ssDNA) fragments of ARS1 identified as the origin of DNA replication (ssARS1-2 and ssARS1-5), but not to the corresponding double-stranded ARS1 (dsDNA) [26]. This specific Sld3–ssDNA association is not affected by CDK phosphorylation. Furthermore, the structures of the Sld7CTD dimer (PDBID:3X38) [22], Sld7NTD–Sld3NTD (PDBID:3X37) [22], Sld3CBD (PDBID:3WI3) [23], MCM DH (6F0L) [27], CMG (PDBIDs:3JC5, 3JC6, 3JC7) [28], CMG-DNA-polɛ (PDBID:7Z13) [29], CMG-DONSON-DNA (PDBID:8Q6O) [30], and so forth, have been determined through crystallography and cryogenic electron microscopy (cryo-EM). Cdc45 belongs to the DHH superfamily of proteins defined by the conserved triad motif DHH (Asp-His-His), and contains a DHH-associated domain (DHHA1: R523–L650) at its C-terminus [31]. Recent single-molecule biochemical assays have reported the stepwise recruitment of multiple Cdc45s to the MCM DH [32]. However, how Sld3–Sld7 recruits Cdc45 onto the MCM for CMG formation to regulate the initiation of DNA replication remains unclear.

Earlier version: In the present study, we aimed to understand how Sld3 recruits Cdc45 to the MCM DH with Sld7 for CMG formation through structure and particle analyses. We determined the structure of S. cerevisiae Sld3CBD–Cdc45 at 2.6 Å resolution and presented the detailed interactions between Sld3 and Cdc45, confirmed through in vitro and in vivo mutant analyses. Compared to the monomer structures, the conformation of Sld3CBD and Cdc45 in the Sld3CBD–Cdc45 complex changed significantly to bind each other. Based on the structural similarity of Cdc45 in Sld3CBD–Cdc45 and CMG, we modelled Sld3CBD–Cdc45–MCM–dsDNA and SCMG–dsDNA (Sld3CBD–CMG–dsDNA) as a snapshot of how helicase CMG forms. The models demonstrated that Sld3CBD, MCM, and GINS bind to different sites on Cdc45, indicating that Sld7–Sld3 could remain at Cdc45–MCM until CMG formation after GINS loading. Consistency between the particle size of Sld7–Sld3ΔC–Cdc45, (Sld3ΔC: M1–K430; truncated the C-terminal domain) as per spectroscopic analysis, and the distance of Sld3CBDs in the Cdc45–MCM dimer suggested that the ternary complex of Sld7–Sld3–Cdc45 forms a dimer off/on the MCM DH for recruiting Cdc45. Furthermore, ssDNA-binding analysis of ARS1 fragments implied that the release of Sld3–Sld7 could be associated with ssARS1 unwound by CMG. Our findings illustrate the recruit–release function of Sld3 in CMG formation and expand our knowledge of the initiation process of DNA replication.

Revised version: The present study aimed to understand how Sld3 recruits Cdc45 to the MCM DH with Sld7 for CMG formation through structure and particle analyses. We determined the structure of S. cerevisiae Sld3CBD–Cdc45 at 2.6 Å resolution and presented the detailed interactions between Sld3 and Cdc45, confirmed through in vitro and in vivo mutant analyses. Compared to the monomer structures, the conformation of Sld3CBD and Cdc45 in the Sld3CBD–Cdc45 complex changed significantly for binding to each other. Based on the structural similarity of Cdc45 in Sld3CBD–Cdc45 and CMG, we modelled Sld3CBD–Cdc45–MCM–dsDNA and SCMG–dsDNA (Sld3CBD–CMG–dsDNA) as a snapshot of how helicase CMG forms. The models demonstrated that Sld3CBD, MCM, and GINS bind to different sites on Cdc45, indicating that Sld7–Sld3 could remain at the Cdc45–MCM until CMG formation after GINS loading. Consistency between the particle size of Sld7–Sld3ΔC–Cdc45 (Sld3ΔC: M1–K430; truncated the C-terminal domain) as per spectroscopic analysis, and the distance of Sld3CBDs in the Cdc45–MCM dimer suggested that the ternary complex of Sld7–Sld3–Cdc45 forms a dimer off/on the MCM DH for recruiting Cdc45. Furthermore, ssDNA-binding analysis of ARS1 fragments suggested that the release of Sld3– Sld7 could be associated with ssARS1, unwound by CMG. Our findings illustrate the recruit– release function of Sld3 in CMG formation, expanding our knowledge of the initiation process of DNA replication.

Results

The overall structure of Sld3CBD and Cdc45 complex

Earlier version: We obtained a recombinant Sld3CBD–Cdc45 complex of S. cerevisiae and determined its crystal structure at 2.6 Å resolution using molecular replacement (Figure 1A, Supplementary Figure 2 A, 3). Only one complex exists within an asymmetric unit. Similar to monomer Sld3CBD (PDBID: 3WI3 with RMSD of 0.50 Å for 181 Cα atoms) [23], Sld3CBD (Y154–P420) in the Sld3CBD–Cdc45 complex is an α-helical structure with two disordered regions of R317–S336 and P364–A369, but with a significant difference in that a disordered part in monomer Sld3CBD was visualized as a C-terminal part of bent long helix α8 (F294–R316; referred to as α8CTP) (Supplementary Figures 3, 4A and 5). Cdc45 (M1–L650) is an α/β structure composed of three β-sheets (anti-parallel: β1-β6-β5-β4-β2-β3, anti-parallel: β7-β8, mixed: β9-β10-β11-β13-β12) surrounded by 21 α-helices (Supplementary Figure 3). Owing to the poor electron density map, the regions D106–K110, D166–K227, S306–V310, and T438–D460 on the molecular surface could not be built. In contrast to the monomer human Cdc45 (huCdc45) and the CMG form (Cdc45 in the yeast CMG complex), more than the N-terminal half of the protruding long helix α7 was disordered in the Sld3CBD–Cdc45 complex (Supplementary Figures 4B and 6).

Revised version: We obtained a recombinant Sld3CBD–Cdc45 complex of S. cerevisiae and determined its crystal structure at 2.6 Å resolution using molecular replacement (Figure 1A, Supplementary Figures 2A and 3). Only one complex molecule Sld3CBD–Cdc45 exists within an asymmetric unit. Similar to the monomeric Sld3CBD (PDBID: 3WI3, with 0.50 Å RMSD for 181 Cα atoms) [23], Sld3CBD (Y154–P420) in the Sld3CBD–Cdc45 complex is an α-helical structure with two disordered regions (R317–S336 and P364–A369). Interestingly, a disordered part in monomeric Sld3CBD was visualized as a C-terminal part of long bent helix α8 (F294–R316; hereafter referred to as α8CTP) (Supplementary Figures 3, 4A and 5). Cdc45 (M1–L650) is an α/β structure composed of three β-sheets (anti-parallel: β1-β6-β5-β4-β2-β3, anti-parallel: β7-β8, mixed: β9-β10-β11-β13-β12) surrounded by 21 α-helices (Supplementary Figure 3). Owing to the poor electron density map, the regions D106–K110, D166–K227, S306–V310, and T438–D460 on the molecular surface could not be built. In contrast to the monomeric human Cdc45 (huCdc45) and the CMG form (Cdc45 in the yeast CMG complex), the N-terminal part of the protruding long helix α7 D219–H231 was disordered in the Sld3CBD–Cdc45 complex (Supplementary Figures 4B and 6).

Sld3CBD–Cdc45 complex structure

(A) Structure of the Sld3CBD–Cdc45 complex. Sld3CBD and Cdc45 are colored in green and magenta, respectively. Cdc45-binding parts α8CTP and α9 are labelled and colored in cyan. The DHHA1 of Cdc45 is labelled and colored in dark magenta. (B) Binding part of Sld3CBD-Cdc45 in different viewing. Dotted squares C, D, and E mark three binding sites, corresponding to the bottom panel. (C) Binding site on the α8CTP of Sld3CBD. (C) (D) Two binding sites involving hydrophobic and hydrogen-bond interactions on the α9 of Sld3CBD, respectively. The interacted residues are depicted by sticks and labelled. The black dotted lines show hydrogen bonds.

Earlier version: Conformation changes in Sld3CBD and Cdc45 for binding each other

Revised version: Conformational changes in Sld3CBD and Cdc45 for binding to each other

Earlier version: Sld3CBD binds to Cdc45 in a way similar to a toothed gear (Figure 1B), with a contact surface of 1808 Å2 accounting for approximately 13.3% and 7.0% of the total surface of Sld3CBD and Cdc45, respectively. Two helices (α8 and α9) of Sld3CBD formed a plier structure and gripped the C-terminal domain DHHA1 (R523–L650) of Cdc45 through hydrophobic interactions and hydrogen bonds. The C-terminal loop (L646–L650) and the C-terminal part of α7 (E232–S242) of Cdc45 sandwiched the helix α8CTP of Sld3CBD. Compared to the isolated form (PDBIDs: 5DGO for huCdc45 [31] and 6CC2 for EhCdc45 [33]) and the CMG form (PDBID: 3JC6 [28]), the Cdc45 in the Sld3CBD–Cdc45 complex changed the conformation of DHHA1 to form pockets on its two sides for binding Sld3CBD α8CTP and α9. The remaining part of Cdc45 (∼K517) retained a structure with RMSD values of 1.29 Å (isolated huCdc45), 1.81 Å (isolated EhCdc45), and 1.295 Å (in CMG) for 243, 251, and 361 Cα atoms, respectively.

Revised version: Sld3CBD binds to Cdc45 in a way similar to a toothed gear (Figure 1B), with a contact surface of 1808 Å2 accounting for approximately 13.3% and 7.0% of the total surface of Sld3CBD and Cdc45, respectively. Two helices (α8 and α9) of Sld3CBD formed a plier structure and gripped the C-terminal domain DHHA1 (R523–L650) of Cdc45 through hydrophobic interactions and hydrogen bonds. The C-terminal loop (L646–L650) and the C-terminal part of α7 (E232–S242) of Cdc45 sandwiched the helix α8CTP of Sld3CBD. Compared to the isolated forms (PDBIDs: 5DGO and 6CC2 for huCdc45 [31] and EhCdc45 [33], respectively) and the CMG form (PDBID: 3JC6 [28]), the Cdc45 in the Sld3CBD–Cdc45 complex changed the conformation of DHHA1 to form pockets on its two sides for binding Sld3CBD α8CTP and α9. The remaining part of Cdc45 (∼K517) retained a structure with RMSD values of 1.29 Å (isolated huCdc45), 1.81 Å (isolated EhCdc45), and 1.295 Å (in CMG) for 243, 251, and 361 Cα atoms, respectively.

Earlier version: The helix α8CTP (F294–R316) of Sld3CBD was surrounded by the C-terminal domain DHHA1 and a C-terminal part of the protruded long helix α7 (E232–S242) of Cdc45 but completely disordered in Sld3CBD alone (PDBID: 3WI3) [23] (Figure 1B and Supplementary Figure 4A). Therefore, the helix α8CTP seems to be an intrinsically disordered segment when Sld3 alone but folds into a helix coupled to the binding partner Cdc45 in the Sld3CBD–Cdc45 complex. This α8CTP is essential for binding to Cdc45, as its deletion inhibited cell growth in a previous study [23]. Furthermore, proline substitution for Cdc45 Ser242 (Cdc45-35), which interacts with L307 and T310 in Sld3CBD α8CTP (Figure 1C), conferred temperature-sensitive growth to yeast cells (Supplementary Figure 7). Although an increase in the number of copies of Sld3–Sld7 could weakly suppress cell growth defects, it did not recover the disrupted interaction.

Revised version: The Sld3CBD helix α8CTP (F294–R316), surrounded by the C-terminal domain DHHA1 and a C-terminal part of the protruded long helix α7 (E232–S242) of Cdc45 (Figure 1B), seems to be an intrinsically disordered segment. When Sld3 is alone, it is disordered but folds into a visualizable helix coupled to the binding partner Cdc45 in the Sld3CBD–Cdc45 complex (Supplementary Figure 4A). Previous studies reported that this α8CTP is essential for binding to Cdc45, as its deletion inhibited cell growth [23]. Furthermore, proline substitution for Cdc45 Ser242 (strain Cdc45-35 [12]), which interacts with L307 and T310 in Sld3CBD α8CTP (Figure 1C), conferred temperature-sensitive growth to yeast cells (Supplementary Figure 7). Although an increase in the number of copies of Sld3–Sld7 could weakly suppress cell growth defects, it did not recover the disrupted interaction.

Earlier version: The α9 (L337–E360) of Sld3CBD is the secondary Cdc45-binding region, which is located in a shallow dent formed by the C-terminal helix and a loop with the following β-strand (K520–Q531) of Cdc45 DHHA1 (Figure 1B, 1D and 1E). Three hydrophobic residues (I352, I355, and L356) in Sld3CBD α9 interact with the C-terminal sheet (L522, L527 and V529) and helix (L641, and L647) of Cdc45 DHHA1 (Figure 1D), while the side chains of two negatively charged residues (D344 and D348) form hydrogen bonds through a water molecule to the main and side chains of R523 on a loop of Cdc45 DHHA1 (Figure 1E). A disordered region, L527–V529, of Cdc45 DHHA1 in the isolated form forms a β-sheet in the Sld3CBD–Cdc45 complex and binds to the I352 and I355 of Sld3CBD α9. We substituted single, double, or ternary positively charged or hydrophilic residues at D344, D348, I352, I355, and L356 of Sld3CBD α9 (D344R/D348R: Sld3-2R, I352Y: Sld3-Y, I352S/I355S/L356S: Sld3-3S, I352E/I355E/L356E: Sld3-3E). Based on the CD spectra, we confirmed that these mutants retained the structural elements of Sld3CBD (Supplementary Figure 8). Double and ternary mutants Sld3-2R, Sld3-3S, and Sld3-3E eliminated Cdc45-binding affinity, whereas the single mutant Sld3-Y seemed to retain a faint interaction with Cdc45 (Figure 2A). We also substituted the single, double, or ternary residues R523, L522, L527, V529, L637, and L641 of Cdc45 on Sld3CBD α9 binding sites (L637S/L641S: Cdc45-IIS, L637E/L641E: Cdc45-IIE, L522S/L527S/V529S: Cdc45-IIIS, L522E/L527E/V529E: Cdc45-IIIE, and Cdc45-A: R523A). All Cdc45 mutants disrupted the binding between Sld3CBD and Cdc45, except for Cdc45-A, as surmised from the Sld3CBD–Cdc45 structure (Figure 2B). Although mutant Cdc45-A has reduced three hydrogen bonds with D344 of Sld3CBD, the remaining hydrogen-bond network maintains contact between Sld3CBD and Cdc45. Furthermore, in vivo genetic studies confirmed the importance of these Sld3 residues. By expressing Sld3-3S, Sld3-3E, and Sld3-2R in Sld3, the variant strains caused complete no growth, whereas the Sld3-Y strain maintains cell growth (Figure 2C), indicating that these residues are required to work together to bind Cdc45 and the loss of the Cdc45-binding affinity of Sld3 inhibits the recruitment of Cdc45 and the subsequent formation of the active replicative helicase CMG for DNA replication.

Revised version: The α9 (L337–E360) of Sld3CBD is the secondary Cdc45-binding region, which is located at a shallow dent formed by the C-terminal helix and a loop with the following β-strand (K520–Q531) of Cdc45 DHHA1 (Figure 1B, 1D and 1E). Three hydrophobic residues (I352, I355, and L356) in Sld3CBD α9 interact with the C-terminal sheet (L522, L527 and V529) and helix (L641 and L647) of Cdc45 DHHA1 (Figure 1D), while the side chains of two negatively charged residues (D344 and D348) form hydrogen bonds through a water molecule to the main and side chains of R523 on a loop of Cdc45 DHHA1 (Figure 1E). A disordered region, L527–V529, of Cdc45 DHHA1 in the isolated form forms a β-sheet in the Sld3CBD–Cdc45 complex and binds to I352 and I355 of Sld3CBD α9. We substituted single, double, or triple positively charged or hydrophilic residues at D344, D348, I352, I355, and L356 of Sld3CBD α9 (D344R/D348R: Sld3-2R, I352Y: Sld3-Y, I352S/I355S/L356S: Sld3-3S, I352E/I355E/L356E: Sld3-3E). Based on the CD spectra, we confirmed that these mutants retained the structural elements of Sld3CBD (Supplementary Figure 8). Double and triple mutants Sld3-2R, Sld3-3S, and Sld3-3E eliminated Cdc45-binding affinity, whereas the single mutant Sld3-Y seemed to retain a faint interaction with Cdc45 (Figure 2A). We also substituted the single, double, or triple residues R523, L522, L527, V529, L637, and L641 of Cdc45 on Sld3CBD α9 binding sites (L637S/L641S: Cdc45-IIS, L637E/L641E: Cdc45-IIE, L522S/L527S/V529S: Cdc45-IIIS, L522E/L527E/V529E: Cdc45-IIIE, and Cdc45-A: R523A). All Cdc45 mutants disrupted the binding between Sld3CBD and Cdc45, except for Cdc45-A, as surmised from the Sld3CBD–Cdc45 structure (Figure 2B). Although mutant Cdc45-A eliminated three hydrogen bonds with D344 of Sld3CBD, the remaining hydrogen-bond network maintains contact between Sld3CBD and Cdc45. Furthermore, in vivo genetic studies confirmed the importance of these Sld3 residues. Expression of Sld3-3S, Sld3-3E, and Sld3-2R in Sld3 caused no growth, while the Sld3-Y strain maintained cell growth (Figure 2C). These results demonstrate that the cooperative action of these residues is essential for Cdc45 binding, and loss of Sld3’s Cdc45-binding affinity inhibits Cdc45 recruitment and subsequent formation of active replicative helicase CMG for DNA replication.

Earlier version: Mutation analysis of interacted residues

Revised version: Mutation analysis of interacting residues

Earlier version: (A) In vitro binding analysis was checked using SDS-PAGE after Ni-affinity chromatography extraction of co-overexpressed Cdc45 with each Sld3 mutant. Sld3 was tagged by His-tag to bind to the column. The labels M, W, Y, 3E, 3S, and 2R are explained on the right. (B) In vitro binding analysis was checked using SDS-PAGE after Ni-affinity chromatography extraction of co-overexpressed Sld3 with each Cdc45 mutant. Sld3 was tagged by His-tag to bind to the column. The labels M, W, A, IIIE, IIE, IIIS, and IIS are explained on the right. (C) In vivo cell growth analysis of yeast cells carrying sld3 mutations. The yeast YYK13 cells carrying SLD3 or its mutant plasmids were streaked onto SD and FOA plates and then incubated at 298 K for 3 days. YYK13 yeast is a mutant lacking the SLD3 gene with added SLD3/sld3 mutant gene (YCplac22 plasmid containing SLD3 or sld3 mutant) that grew on SD and FOA plates. The empty plasmids were used as negative control (NC). Mutations in Sld3-Y, Sld3-3E, Sld3-3S, and Sld3-2R are the same as those in A.

Revised version: (A) In vitro binding analysis was checked using SDS-PAGE after Ni-affinity chromatography extraction of co-overexpressed Cdc45 with each Sld3 mutant. Sld3 was tagged by His-tag to bind to the column. The labels M, W, Y, 3E, 3S, and 2R are explained on the right. (B) In vitro binding analysis was checked using SDS-PAGE after Ni-affinity chromatography extraction of co-overexpressed Sld3 with each Cdc45 mutant. Sld3 was tagged by His-tag to bind to the column. The labels M, W, A, IIIE, IIE, IIIS, and IIS are explained on the right. (C) In vivo cell growth analysis of yeast cells carrying sld3 mutations. The yeast YYK13 cells carrying SLD3 or its mutant plasmids were streaked onto SD and FOA plates and then incubated at 298 K for 3 days. YYK13 yeast is a mutant lacking the SLD3 gene with added SLD3/sld3 mutant gene (YCplac22 plasmid containing SLD3 or sld3 mutant) that grew on SD and FOA plates. The empty plasmids were used as a negative control (NC). Mutations in Sld3-Y, Sld3-3E, Sld3-3S, and Sld3-2R are the same as those in A.

Earlier version: In comparison with Cdc45 alone (huCdc45) or CMG-form (in CMG complex), domain DHHA1 of Cdc45 changed conformation significantly for binding to Sld3CBD (Supplementary Figure 4C). The loop I595–N604 in Cdc45 DHHA1 changed conformation to interact with Sld3CBD α8CTP. Subsequent the following helix α19 (F605–E615) rotated the C-terminus by 25 degrees, which altered the conformation of the next two β-strands (β12 and β13) in the mixed β-sheet (β9-β10-β11-β13-β12) (Supplementary Figure 3 and 4C), allowing Sld3CBD α8CTP to enter the binding pocket. Interestingly, the Sld3CBD-Cdc45 structure showed that the Sld3CBD binding site of Cdc45 completely differed from that of GINS and MCM bound to Cdc45, suggesting that the Sld3CBD, Cdc45 and GINS could bind to MCM together (Supplementary Figure 9A). Furthermore, we conducted a mutation analysis on two Cdc45 residues involved in binding to MCM (Cdc45 G367D) and GINS (Cdc45 W481R) [34], and found that these mutations did not disrupt the binding of Sld3CBD-Cdc45 complex (Supplementary Figure 9B).

Revised version: In comparison with Cdc45 alone (huCdc45) or CMG-form (in CMG complex), domain DHHA1 of Cdc45 changed conformation significantly for binding to Sld3CBD (Supplementary Figure 4C). The loop I595–N604 in Cdc45 DHHA1 changed conformation to interact with α8CTP of Sld3CBD. Subsequently, the helix α19 (F605–E615) rotated the C-terminus by 25 degrees, which altered the conformation of the next two β-strands (β12 and β13) in the mixed β-sheet (β9-β10-β11-β13-β12) (Supplementary Figures 3 and 4C), allowing Sld3CBD α8CTP to enter the binding pocket. Interestingly, the Sld3CBD-Cdc45 structure shows that the Sld3CBD binding site of Cdc45 is distinct from the binding site of Cdc45 with GINS or MCM, suggesting that the Sld3CBD, Cdc45 and GINS could bind to MCM without steric clash (Supplementary Figure 9A). Furthermore, we conducted a mutation analysis on two Cdc45 residues involved in binding to MCM (Cdc45 G367D) and GINS (Cdc45 W481R) [34], respectively, and found that these mutations did not disrupt the Sld3CBD-Cdc45 complex (Supplementary Figure 9B).

Cdc45 recruitment to MCM DH by Sld3 with partner Sld7

Earlier version: Except for the Sld3 binding region DHHA1, the N-terminal domain of Cdc45 (Cdc45NTD) maintained a structure similar to that of the monomer, Sld3CBD–Cdc45 complex, and CMG complex. Therefore, we modelled Sld3CBD–Cdc45–MCM–dsDNA and SCMG–dsDNA (Figure 3) by superposing Cdc45NTD structures between Sld3CBD–Cdc45 and CMG dimers respectively, which were modelled by superposing Mcm2–7 structures between CMG (PDBID: 3JC6) [28] and MCM–dsDNA (PDBID: 5BK4) [35]. In the models, two Sld3CBDs are located in each monomer of the Cdc45–MCM dimer over 230 Å apart (Figure 3A).

Revised version: Except for the Sld3 binding region DHHA1, the N-terminal domain of Cdc45 (Cdc45NTD) retained a structure similar to that in the monomer, Sld3CBD–Cdc45 complex and CMG complex. Therefore, we modelled Sld3CBD–Cdc45–MCM–dsDNA and SCMG–dsDNA (Figure 3) by superposing Cdc45NTD structures between Sld3CBD–Cdc45 and CMG dimers respectively, which were modelled by superposing Mcm2–7 structures between CMG (PDBID: 3JC6) [28] and MCM–dsDNA (PDBID: 5BK4) [35]. In the models, two Sld3CBDs are located in each monomer of the Cdc45–MCM dimer over 230 Å apart (Figure 3A).

Ribbon models of complexes in dimer form and particle analysis

Earlier version: (A) Sld3CBD–Cdc45–MCM–dsDNA complex. Mcm2, 5, 4, and 6 subunits are colored in cyan, blue, marine, and light blue, respectively. Subunits Mcm3 and Mcm7 are colored in gray. Green and pink are used to color Sld3CBD and Cdc45, respectively. dsDNA is represented by a dark-orange stick. (B) Sld3CBD–Sld7–Cdc45 dimer before associating with the MCM DH. Sld3 and Cdc45 are shown in the same color as they are in A, while Sld7 is colored in orange. The two phosphorylated Sld3 residues are depicted as yellow balls. Particle analysis of Sld7–Sld3ΔC–Cdc45 through dynamic light scattering is shown on the bottom panel. The average peak size of the particle size distribution of the Sld7–Sld3ΔC–Cdc45 complex was estimated to be 23.2 Å in diameter. The measurement was carried out independently four times (Supplementary Figure 9). (C) SCMG–dsDNA complex. GINS is shown in yellow and the remainder are colored identically to those in A and B.

Revised version: (A) Sld3CBD–Cdc45–MCM–dsDNA complex. Mcm2, 5, 4, and 6 subunits are colored in cyan, blue, marine, and light blue, respectively. Subunits Mcm3 and Mcm7 are colored in gray. Green and pink are used to color Sld3CBD and Cdc45, respectively. dsDNA is represented by a dark-orange stick. (B) Sld3CBD–Sld7–Cdc45 dimer before associating with the MCM DH. Sld3 and Cdc45 are shown in the same color as they are in A, while Sld7 is colored in orange. The two phosphorylated Sld3 residues are depicted as yellow balls. Particle analysis of Sld7–Sld3ΔC– Cdc45 through dynamic light scattering is shown on the bottom panel. The average peak size of the particle size distribution of the Sld7–Sld3ΔC–Cdc45 complex was estimated to be 232 Å in diameter. The measurement was carried out independently four times (Supplementary Figure 9). (C) SCMG–dsDNA complex. GINS is shown in yellow, and the remainder are colored identically to those in A and B.

Earlier version: To investigate how Sld7–Sld3 brings Cdc45 to MCM DH, we attempted particle analysis for the Sld7–Sld3–Cdc45 complex using the purified recombinant Sld7–Sld3ΔC–Cdc45. The approximate molecular weight of Sld7–Sld3ΔC–Cdc45 was estimated to be >400 kDa according to a weight calibration of size-exclusion chromatography (Supplementary Figure 2B). Given that the molecular weight calculated from their amino acid sequences was 158 kDa, the purified complex should be a dimer. Considering the Sld7–Sld3ΔC–Cdc45 dimer should be a long rod-sharped molecule, the estimated value from SEC could be larger than theoretical values. Subsequently, using dynamic light scattering, the particle size (hydrodynamic diameter) of the tripartite complex was estimated to be around 232 Å (Figure 3B, Supplementary Figure 10), which is consistent with the distance of Sld3CBDs in the model of Sld3CBD–Cdc45–MCM dimer. Furthermore, considering that the domains of Sld7 (NTD: Sld3NTD-binding, CTD: self-dimerization) and Sld3 (NTD: Sld7NTD-binding, CBD: Cdc45-binding, CTD: MCM-binding) function independently [22,23], we estimated a dimer model of Sld7–Sld3–Cdc45, as shown in Figure 3B. Because the domains of Sld7–Sld3–Cdc45 are linked by long loops, the dimer of Sld7–Sld3–Cdc45 is flexible with long-shape.

Revised version: To investigate how Sld7–Sld3 brings Cdc45 to MCM DH, we attempted particle analysis for the Sld7–Sld3–Cdc45 complex using the purified recombinant Sld7–Sld3ΔC–Cdc45. The approximate molecular weight of Sld7–Sld3ΔC–Cdc45 was estimated to be >400 kDa according to a weight calibration of size-exclusion chromatography (Supplementary Figure 2B). Given that the molecular weight calculated from its amino acid sequences was 158 kDa, the purified complex should be a dimer. Considering the Sld7–Sld3ΔC–Cdc45 dimer should be a long, rod-shaped molecule, the estimated value from SEC could be larger than the theoretical values. Subsequently, using dynamic light scattering, the particle size (hydrodynamic diameter) of the tripartite complex was estimated to be around 232 Å (Figure 3B, Supplementary Figure 10), which is consistent with the distance of Sld3CBDs in the model of Sld3CBD–Cdc45–MCM dimer. To further validate the SEC and DLS results, we performed size-exclusion chromatography coupled with small-angle X-ray scattering (SEC-SAXS), which suggested a molecular weight of 370 - 420 kDa, and an Rg > 85 Å (Supplementary Figure 11). Considering that the domains of Sld7 (NTD: Sld3NTD-binding, CTD: self-dimerization) and Sld3 (NTD: Sld7NTD-binding, CBD: Cdc45-binding, CTD: MCM-binding) function independently [22,23], we estimated a dimer model of Sld7–Sld3ΔC–Cdc45, as shown in Figure 3B. Because the domains of Sld7–Sld3–Cdc45 are linked by long loops, the dimer forms a long shape with high flexibility.

Earlier version: Our SCMG–dsDNA model demonstrated that Sld3CBD neighbors Mcm2 and binds to Cdc45 on the opposite side of GINS binding (Figure 3C), indicating that Sld7–Sld3 bound to the Cdc45–MCM dimer did not contact GINS and could be remained for CMG formation. A recent study also reported a structure of the DONSON (Sld2 homolog) dimer with CMG (8Q6O), showing that the DONSON dimer delivers GINS to MCM and reconfigures the MCM motors in the double CMG [30]. The DONSON dimer is loading at the GINS site on MCM, which is the opposite of Sld3CBD in the SCMG-dsDNA model. Interestingly, in the Sld3CBD–Cdc45–MCM–dsDNA and SCMG–dsDNA model, we found that a part (S358-D383) of Sld3CBD containing the C-terminal of α9, a disordered fragment, and α10 was in contact distance with Mcm2CTD. In addition, the Sld3CBD-bound Cdc45 DHHA1 appeared to become close to Mcm2CTD. In contrast, the Cdc45 DHHA1 does not contact with Mcm2–7 or GINS in the CMG structure (Supplementary Figure 11). The conformational change in Cdc45 DHHA1 not only facilitates binding with Sld3CBD, but could also cause contact with Mcm2 without affecting the interaction between Cdc45 and GINS.

Revised version: Our SCMG–dsDNA model demonstrated that Sld3CBD neighbors Mcm2 and binds to Cdc45 on the opposite side of GINS binding (Figure 3C), indicating that Sld7–Sld3 bound to the Cdc45– MCM dimer does not contact GINS and could remain until CMG formation. A recent study also reported a structure of the DONSON (Sld2 homolog) dimer with CMG (8Q6O), showing that the DONSON dimer delivers GINS to MCM and reconfigures the MCM motors in the double CMG [30]. The DONSON dimer is loaded at the GINS site on MCM, which is the opposite of Sld3CBD in the SCMG-dsDNA model. Interestingly, in the Sld3CBD–Cdc45–MCM–dsDNA and SCMG– dsDNA models, we found that a part (S358-D383) of Sld3CBD, containing the C-terminal of α9, a disordered fragment, and α10, was in contact distance with Mcm2CTD. In addition, the Sld3CBD-bound Cdc45 DHHA1 appeared to be close to Mcm2CTD. In contrast, the Cdc45 DHHA1 does not contact Mcm2–7 or GINS in the CMG structure (Supplementary Figure 12). The conformational change in Cdc45 DHHA1 not only facilitates binding with Sld3CBD, but could also lead to contact with Mcm2 without affecting the interaction between Cdc45 and GINS.

ssDNA binding affinity of Sld3 depended on complex formation with Cdc45 and Sld7

Earlier version: Previous studies showed that Sld3 binds directly to single-strand DNA fragments (ssARS1-2 and ssARS1-5) of ARS1 identified as an origin of DNA replication [26]. ARS1 was divided into three 80-bp segments: ARS1-12, ARS1-34, and ARS1-56, and these dsDNA segments could unwind into six signal-stranded DNA fragments with 80b length: ssARS1-1, 2, 3, 4, 5, and 6 (Supplementary Figure 12A). Given that Sld3 binds to Sld7 and Cdc45 on the MCM–DNA complex during CMG formation, we investigated whether Sld7 and Cdc45 affect the ssDNA-binding affinity of Sld3. Therefore, we performed an ssDNA binding assay using the Sld3CBD, Sld3CBD–Cdc45, Sld7–Sld3ΔC–Cdc45 and Sld7–Sld3ΔC (Sld3ΔC: binding part M1–K430 to Sld7 and Cdc45) complexes against different regions of ssARS1 for comparison. Cdc45 binds tighter to long ssDNA (>60 bases) with a litter sequence specificity [36]. Therefore, we fragmented each ssARS1 into fragments of 40 bases to prevent this nonspecific interaction (Figure 4A and Supplementary Figures 12B).

Revised version: Previous studies showed that Sld3 binds directly to single-strand DNA fragments (ssARS1-2 and ssARS1-5) of ARS1 identified as an origin of DNA replication [26]. ARS1 was divided into three 80-bp segments: ARS1-12, ARS1-34, and ARS1-56. These dsDNA segments could unwind into six single-stranded DNA fragments of 80 nucleotides (nt) in length: ssARS1-1, 2, 3, 4, 5, and 6 (Supplementary Figure 12A). Given that Sld3 binds to Sld7 and Cdc45 on the MCM–DNA complex during CMG formation, we investigated whether Sld7 and Cdc45 affect the ssDNA-binding affinity of Sld3. Therefore, we performed an ssDNA binding assay using the Sld3CBD, Sld3CBD–Cdc45, Sld7–Sld3ΔC–Cdc45 and Sld7–Sld3ΔC (Sld3ΔC: M1–K430, domains for binding Sld7 and Cdc45) complexes against different regions of ssARS1 for comparison. Due to limitations in protein overexpression, we utilized Sld7–Sld3ΔC–Cdc45 and Sld7–Sld3ΔC from K. marxianus (same family as S. cerevisiae). Moreover, Cdc45 exhibits higher affinity for >60nt ssDNA compared with shorter ssDNA [36]. Therefore, we fragmented each ssARS1 into fragments of 40nt to prevent such nonspecific binding (Figure 4A and Supplementary Figure 13).

Electrophoresis mobility shift assay (EMSA) of ssDNA binding to Sld3 and its complexes with Sld7 and Cdc45

Earlier version: (A) Schematic of ssARS1-1–ssARS1-6. ARS1 is identified as an origin of DNA replication, and divided into three 80-bp segments: ARS1-12, ARS1-34, and ARS1-56. These dsDNA segments could unwind into six signal-stranded DNA fragments with 80b length: ssARS1-1, 2, 3, 4, 5, and 6. The important elements of A (ARS consensus sequence), B1, B2, and B3 for unwinding are marked [37]. We divided ssARS1-2 and ssARS1-5 fragments (blue squares) into 40-base lengths for EMSA. (B) ssDNAs were visualized using SYBR safe on polyacrylamide gels. In the presence of ssDNA fragments, Sld3CBD, Sld3CBD–Cdc45, Sld7–Sld3ΔC–Cdc45 and Sld7–Sld3ΔC were incubated with molecular mass-related concentrations. The molecular ratio of ssDNA to protein in lanes 1, 2, 3, 4, and 5 was 1:0, 1:1, 1:2, 1:4, and 0:1, respectively. The controls for ssDNA and protein are lanes 1 and 5, respectively. The disappearance of the ssDNA band in lanes 2–5 indicates that the protein (Sld3CBD, Sld7–Sld3ΔC and Sld7–Sld3ΔC–Cdc45) binds with high affinity.

Revised version: (A) Schematic of ssARS1-1–ssARS1-6. ARS1 is identified as an origin of DNA replication and divided into three 80-bp segments: ARS1-12, ARS1-34, and ARS1-56. These dsDNA segments could unwind into six single-stranded DNA fragments with 80nt length: ssARS1-1, 2, 3, 4, 5, and 6. The important elements of A (ARS consensus sequence), B1, B2, and B3 for unwinding are marked [37]. We divided ssARS1-2 and ssARS1-5 fragments (blue squares) into 40-base lengths for EMSA. (B) ssDNAs were visualized using Fast Blast DNA stain on polyacrylamide gels. In the presence of ssDNA fragments, Sld3CBD, Sld3CBD–Cdc45, Sld7–Sld3ΔC–Cdc45 and Sld7– Sld3ΔC were incubated with molecular mass-related concentrations. The molecular ratio of ssDNA to protein in lanes 1, 2, 3, 4, and 5 was 1:0, 1:0.5, 1:1, 1:2, and 0:1, respectively. The controls for ssDNA and protein are lanes 1 and 5, respectively. No binding ssDNA group (ssARS1-3_1) is shown at the bottom as a negative control (NC). The overall views of the EMSA results are shown in Supplementary Figure 14. (C) The integral grayscale of the ssDNA bands was calculated and compared to the average of the ssDNA control band to determine the residual levels, showing differences in binding affinity. By three ratios, Sld3CBD-Cdc45 demonstrated a significantly ssDNA residual level (t-test, ****: P<0.0001) compared to other samples, indicating low binding affinity to ssDNA.

Earlier version: To investigate the specificity of the ssDNA binding affinity of Sld3, we employed an electrophoretic mobility shift assay (EMSA) using non-denaturing PAGE (native-PAGE) with ssDNAs or proteins alone as controls. For the Sld3CBD and ssDNA mixtures at a molar ratio of 1:1, the band corresponding to ssDNA disappeared, indicating a strong binding affinity (Figure 4B). The ssDNA band remained (Figure 4B) and new bands corresponding to the ssDNA–protein complex appeared in CBB staining PAGE (Supplementary Figures 13) when the Sld3CBD–Cdc45 complex was mixed with ssDNA at the same ratio, indicating that the binding affinity of Sld3CBD–Cdc45 for ssDNA was lower than that of Sld3CBD alone. Additionally, the disappearance of ssDNA bands and the appearance of new bands corresponding to the ssDNA–protein complex correlated with an increase in the concentration of Sld3CBD–Cdc45 in the mixture.

Revised version: To investigate the specificity of the ssDNA binding affinity of Sld3, we employed an electrophoretic mobility shift assay (EMSA) using non-denaturing PAGE (native-PAGE) with ssDNAs or proteins alone as controls (Figure 4B and Supplementary Figure 14). Additionally, an ssDNA (ssARS1-3_1) with no binding affinity to protein samples was used as a negative control (NC), where no ssDNA band disappeared and no new ssDNA band appeared. For the Sld3CBD and ssDNA mixtures at a molar ratio of 1:1, the band corresponding to ssDNA disappeared, indicating a binding affinity (Figure 4C). The ssDNA band remained when the Sld3CBD–Cdc45 complex was mixed with ssDNA at the same ratio, indicating that the binding affinity of Sld3CBD– Cdc45 for ssDNA was lower than that of Sld3CBD alone (Figure 4C). Additionally, the decrease of ssDNA bands correlated with an increase in the concentration of Sld3CBD–Cdc45 in the mixture.

Earlier version: In the presence of the Sld7–Sld3ΔC–Cdc45 complex, the band corresponding to ssDNA disappeared in the equimolar ratio of protein and ssDNA, similar to that observed for Sld3CBD (Figure 4B and Supplementary Figures 13). This result demonstrates that the presence of Sld7 and Sld3NTD could increase the ssDNA-binding affinity to a level comparable to that of Sld3CBD. Also, Sld7–Sld3ΔC showed a ssDNA-binding affinity similar to that of Sld7–Sld3ΔC–Cdc45, implying that ssDNA-binding of Sld7–Sld3ΔC is independent of Cdc45. Furthermore, the results revealed no significant difference in binding affinity between the 40-base fragments of ssARS1-2-1, 2, 3 and ssARS1-5-1, 2, 3 for Sld3CBD, Sld3CBD–Cdc45, Sld7–Sld3ΔC–Cdc45, and Sld7–Sld3ΔC, indicating that there is no stronger binding-region specific to ssARS1-2 or ssARS1-5 fragments. For sequence specificity, we also analyzed other fragments of ssARS1-1, 3, 4, and 6 and dsARS1, and all of them showed no binding (Supplementary Figure 14). Such binding experiments can serve as a negative control, confirming that the absence or presence of ssDNA bands is contingent on protein binding.

Revised version: In the presence of the Sld7–Sld3ΔC–Cdc45 complex, the band corresponding to ssDNA disappeared under the equimolar ratio of protein and ssDNA, similar to that observed for Sld3CBD (Figure 4C and Supplementary Figure 14A). In this case, smeared bands appeared in the high molecular weight part of ssDNA-staining PAGE (Figure 4B, Supplementary Figure 14). The positions of smeared ssDNA bonds correspond to those of protein in the protein-stain pages, indicating that ssARS1 were complexed with proteins. This result demonstrates that the presence of Sld7 and Sld3NTD could increase the ssDNA-binding affinity to a level comparable to that of Sld3CBD. Also, Sld7–Sld3ΔC showed the ssDNA-binding affinity similar to that of Sld7–Sld3ΔC– Cdc45, implying that the ssDNA-binding of Sld7–Sld3ΔC is independent of Cdc45. Furthermore, the results revealed no significant difference in binding affinity between the 40-base fragments of ssARS1-2-1, 2, 3 and ssARS1-5-1, 2, 3 for Sld3CBD, Sld3CBD–Cdc45, Sld7–Sld3ΔC–Cdc45, and Sld7–Sld3ΔC, indicating that there is no stronger binding-region specific to ssARS1-2 or ssARS1-5 fragments. For sequence specificity, we also analyzed other fragments (ssARS1-1, 3, 4, and 6, and dsARS1), and all of them showed no binding (Supplementary Figure 15).

Earlier version: The surface charge of the Sld3CBD–Cdc45 structure, shows that Cdc45 covers the main positive charge region of Sld3CBD α8CTP (Supplementary Figure 15A), which may weaken the binding affinity of Sld3CBD–Cdc45 to ssDNA. Conversely, on the Sld3–Sld7 structure, there is a large positive charge area strip on Sld7NTD (Supplementary Figure 15B). Considering that ssARS1 is unwound from dsARS1 by the activated helicase CMG complex formed after loading Cdc45 and GINS, Sld3–Sld7 having a stronger ssARS1-binding affinity may provide an advantage for the dissociation of Sld7–Sld3 from the CMG complex.

Revised version: The surface charge of the Sld3CBD–Cdc45 structure shows that Cdc45 covers the main positive charge region of Sld3CBD α8CTP (Supplementary Figure 16A), which may weaken the binding affinity of Sld3CBD–Cdc45 to ssDNA. Conversely, on the Sld3–Sld7 structure, there is a large positive charge area strip on Sld7NTD (Supplementary Figure 16B). Considering that ssARS1 is unwound from dsARS1 by the activated helicase CMG complex formed after loading Cdc45 and GINS, the stronger binding affinity of Sld3–Sld7 may provide an advantage for the dissociation of Sld7–Sld3 from the CMG complex.

Discussion

Earlier version: As a central regulator of helicase CMG formation, Sld3 has attracted interest in studies aiming to understand the initiation of DNA replication. Owing to the lack of structural information on Sld3 complexed with Cdc45 or MCM, how Sld3 regulates helicase activation with other factors remains unknown. Here, we present the structure of the Sld3CBD–Cdc45 complex and particle size of the Sld7–Sld3–Cdc45 complex to examine the detailed mechanisms underlying CMG formation.

Revised version: As a central regulator of helicase CMG formation, Sld3 has attracted interest in studies aiming to understand the initiation of DNA replication. Owing to the lack of structural information on Sld3 complexed with Cdc45 or MCM, how Sld3 regulates the formation of activated helicase CMG with other factors remains unknown. Here, we present the structure of the Sld3CBD–Cdc45 complex and particle size of the Sld7–Sld3–Cdc45 complex to examine the molecular mechanisms underlying CMG formation.

Earlier version: Sld3 is highly conserved in eukaryotes, whereas its functional homolog/counterpart in metazoans Treslin (also known as Ticrr) has a distinct size and sequence, except for Sld3CBD. Sequence alignment of Sld3CBD among Sld3 and Treslin revealed that all Cdc45-binding residues in α8 and α9 identified in our study were almost conserved or exhibited conserved changes (Supplementary Figures 4 and 5), indicating that these regions provide a similar mode of interaction between Sld3CBD and Cdc45 in the regulation of metazoan DNA replication. Therefore, we hypothesize that they load Cdc45, as observed in yeast Sld3 and Sld7.

Revised version: Sld3 exhibits high conservation across eukaryotes, whereas its functional ortholog in metazoans Treslin (also known as Ticrr), has a distinct size and sequence, except for the Cdc45-binding domain (Sld3CBD). Sequence alignment of Sld3CBD among Sld3 and Treslin revealed that all Cdc45-binding residues in α8 and α9 identified in our study were almost conserved or exhibited conserved changes (Supplementary Figures 4 and 5). This conservation suggests that these regions provide a similar interaction manner between Sld3CBD and Cdc45 in the regulation of metazoan DNA replication. Therefore, we hypothesize that Treslin may load Cdc45 as observed in yeast Sld3 and Sld7.

Earlier version: Sld3CBD and Cdc45 significantly change their conformation to bind to each other in the Sld3CBD–Cdc45 complex compared with their conformation as monomers. Interestingly, the conformational changes in Cdc45 DHHA1 upon binding to Sld3CBD also caused contact with Mcm2NTD in the Sld3CBD–Cdc45–MCM–dsDNA and SCMG–dsDNA models, whereas DHHA1 interacted with neither MCM nor GINS in the CMG structure. Taken together, Sld3CBD contacts Mcm2NTD in the Sld3CBD–Cdc45–MCM–dsDNA and SCMG–dsDNA models and then seems to play a guiding role in helping Cdc45 bind to MCM at the right position. Sld3CBDs in each monomer of the Sld3CBD–Cdc45–MCM dimer were located at a distance of more than 230 Å, which is consistent with the results of a particle size analysis of the Sld7–Sld3ΔC–Cdc45 complex off MCM DH in solution. Therefore, we propose that a binding manner of Sld7–Sld3–Cdc45 in a flexible long-sharped dimer Cdc45–Sld3–(Sld7)2–Sld3–Cdc45 off/on MCM DH is advantageous for efficiently recruiting two Cdc45 molecules to an MCM DH, consequently leading to the formation of a pair of CMG helicases.

Revised version: By structural comparison, we found that Sld3CBD and Cdc45 changed their conformations to bind to each other. The conformational changes in Cdc45 DHHA1 upon binding to Sld3CBD also caused the contact between Cdc45 and Mcm2NTD in the Sld3CBD–Cdc45–MCM–dsDNA and SCMG–dsDNA models, whereas DHHA1 interacted with neither MCM nor GINS in the CMG structure. Taking the structural information together, Sld3 seems to play a guiding role in helping Cdc45 bind to MCM at the right position. Furthermore, Sld3CBDs in each monomer of the Sld3CBD–Cdc45–MCM dimer were located at a distance of more than 230 Å, which is consistent with the results of a particle size analysis of the Sld7–Sld3ΔC–Cdc45 complex off MCM DH in solution. Therefore, we propose that a binding manner of Sld7–Sld3–Cdc45 in a flexible long-shaped dimer Cdc45–Sld3–(Sld7)2–Sld3–Cdc45 off/on MCM DH is advantageous for efficiently recruiting two Cdc45 molecules to an MCM DH, consequently leading to the formation of a pair of CMG helicases.

Earlier version: In the SCMG–dsDNA complex model, Cdc45 bound to Sld3CBD, MCM, and GINS on different sides (with contact surfaces of 6.7, 4.8, and 5.1% of the total Cdc45 surface, respectively). Our structure of Sld3CBD-Cdc45 and model shows that these bindings occur at different sites of Cdc45, suggesting that Sld3CBD, Cdc45 and GINS could bind to MCM together, and do not directly interfere with each other. The competition of Sld3 and GINS for binding to Cdc45 or Cdc45-MCM reported by using in vitro binding experiments [26], may be caused by the conformation change of Cdc45 DHHA1 or without other initiation factors. In particular, Sld3 and GINS bind to Cdc45 and MCM rings at opposite positions (Mcm2–4–6 and Mcm5–3–7 vs. Sld3 and GINS, respectively), suggesting that the GINS-recruitment protein should cross a long-distance in an MCM monomer or MCM DH to recognize the phosphorylation site (T600 and S622) of Sld3. Furthermore, our SCMG–dsDNA model revealed that Sld3CBD on CMG appears to contact an N-terminal helix of Mcm2CTD and Sld3CTD could extend to bind to Mcm4NTD through interaction with the NTDs of Mcm2, and 6, rather than the CTD tier (Supplementary Figure 16), indicating that the binding of Sld3–Sld7 to MCM does not appear to affect the ability of AAA+ motors to open and close the MCM ring with a gap between Mcm2CTD and Mcm5CTD [27, 28, 35, 38]. Taking our findings and those of previous studies together, we propose a detailed process for helicase CMG formation from inactive MCM, as depicted in Figure 5 A–C, in which Sld7–Sld3 brings Cdc45 into the MCM as a Sld7–Sld3–Cdc45 dimer (Cdc45–Sld3–[Sld7]2–Sld3–Cdc45), which remains until GINS loading.

Revised version: In the SCMG–dsDNA complex model, Cdc45 bound to Sld3CBD, MCM, and GINS on different sides (with contact surfaces of 6.7, 4.8, and 5.1% of the total Cdc45 surface, respectively). Our structure of Sld3CBD-Cdc45 and models show that these bindings occur at distinct sites on Cdc45, suggesting that Sld3CBD, Cdc45 and GINS could bind to MCM together without steric clash. The competition between Sld3 and GINS for binding to Cdc45 or Cdc45-MCM (by mixing them in vitro) reported by Bruck et al. [26] may be caused by the conformational change of Cdc45 DHHA1 or the lack of other auxiliary initiation factors, indicating that activated CMG formation requires regulation. In particularly, Sld3 and GINS bind to opposite positions of Cdc45 and MCM ring (Mcm2–4–6 vs Sld3 and Mcm5–3–7 vs GINS), suggesting that the GINS-recruitment protein should cross a long-distance in an MCM monomer or MCM DH to access the phosphorylation site (T600 and S622) of Sld3. Furthermore, our SCMG–dsDNA model revealed that Sld3CBD on CMG appears to contact an N-terminal helix of Mcm2CTD, while Sld3CTD may extend to bind to Mcm4NTD through interaction with the Mcm2NTD and Mcm6NTD (Supplementary Figure 17). These findings suggest that the Sld3–Sld7 binding to MCM does not interfere with the AAA+ motors’ ability to regulate MCM ring dynamics during its opening/closing via the gap between Mcm2CTD and Mcm5CTD [27, 28, 35, 38]. Taking our findings and those of previous studies together, we propose a detailed process for helicase CMG formation from inactive MCM, as depicted in Figure 5 A–C: Sld7–Sld3 brings Cdc45 onto MCM as a Sld7–Sld3–Cdc45 dimer (Cdc45–Sld3–[Sld7]2–Sld3–Cdc45), and remains until GINS loading.

Proposal for CMG formation with Sld7–Sld3

Earlier version: (A) Phosphorylation of Mcm2,4,6 by DDK after MCM DH was loaded on dsDNA at the replication origin. (B) Cdcd45 recruitment to MCM DH by Cdc45–Sld3–[Sld7]2–Sld3–Cdcd45. (C) After CDK-mediated phosphorylation of Sld3CTD in Cdc45–MCM, Dpb11–Sld2 recruits GINS and polε to Sld7–Sld3-Cdc45–MCM to form an active helicase CMG. (D) Unwinding of dsDNA by CMG with MCM DH separation and MCM ring opening. Sld3 and other factors are released upon binding to ssDNA. (E) Each CMG unwinds the dsDNA in two directions, initiating DNA replication.

Revised version: (A) Phosphorylation of Mcm2,4,6 by DDK after MCM DH was loaded on dsDNA at the replication origin. (B) Cdc45 recruitment to MCM DH by Cdc45–Sld3–[Sld7]2–Sld3– Cdc45. (C) After CDK-mediated phosphorylation of Sld3CTD in Cdc45–MCM, Dpb11– Sld2 recruits GINS and polε to Sld7–Sld3-Cdc45–MCM to form an active helicase CMG. (D) Unwinding of dsDNA by CMG with MCM DH separation and MCM ring opening. Sld3 and other factors are released upon binding to ssDNA. (E) Each CMG unwinds the dsDNA in two directions, initiating DNA replication.

Earlier version: The following inquiry concerns the dissociation of Sld3 and other factors. Interestingly, the mutant analysis showed that disrupting just one binding site between Sld3CBD and Cdc45 could cause the dissociation of Sld3CBD and Cdc45, indicating that the binding between Sld3CBD and Cdc45 could be broken easily. Our binding analysis of ssARS1 fragments to Sld3CBD, Sld3CBD–Cdc45, and Sld7–Sld3ΔC–Cdc45 demonstrated the sequence specificity to ssARS1-2 and ssARS1-5 fragments, not others including dsARS1. Considering that ssARS1-2 and ssARS1-5 are on both sides of MCM DH-bound dsDNA at replication origin [39], the origin unwinding by CMG generates ssDNA and further sequesters the Sld7-Sld3 complex onto ssDNA to remove Sld7-Sld3 from CMG. As a bridge protein, Sld3 recruits Cdc45 to MCM, and phosphorylated Sld3 regulates the recruitment of GINS [40]. Finally, the release of Sld3 and Sld7 from CMG could be associated with the binding of ssARS1 unwound by CMG and may also be related to the dissociation of Dpb11–Sld2 from CMG (Figure 5 D, E) [29, 41]. Further structural investigation of Sld3–Sld7–Cdc45–MCM in the GINS-recruiting stage is required to validate the complete CMG formation process.

Revised version: The following inquiry concerns the dissociation of Sld3 and other factors. Interestingly, the mutant analysis demonstrates that disrupting a single binding site between Sld3CBD and Cdc45 suffices to dissociate Sld3CBD and Cdc45, indicating that a functionally critical binding between Sld3CBD and Cdc45 can be broken easily. Furthermore, our binding analysis of ssARS1 fragments to Sld3CBD, Sld3CBD–Cdc45, Sld7–Sld3ΔC–Cdc45 and Sld7–Sld3ΔC showed the sequence specificity to ssARS1-2 and ssARS1-5 fragments, not others of dsARS1. Considering that ssARS1-2 and ssARS1-5 are on both sides of MCM DH-bound dsDNA at replication origin [39], the origin unwinding by CMG generates ssDNA and further sequesters the Sld7-Sld3 complex onto ssDNA to remove Sld7-Sld3 from CMG. As a bridge protein, Sld3 recruits Cdc45 to MCM, and its next phosphorylated state regulates the subsequent recruitment of GINS loading with Dpb11–Sld2 [40]. Thus, the release of Sld3 and Sld7 from CMG could be associated with unwound ssARS1 and may also be related to the dissociation of Dpb11–Sld2 from CMG (Figure 5D, E) [29, 41]. Furtherly, our proposals require a visualization of the Sld3–Sld7–Cdc45–MCM complex structure during GINS recruitment to establish the complete CMG formation process.

Earlier version: In conclusion, our structural and biochemical studies of Sld3CBD–Cdc45 revealed a detailed process of CMG formation and the subsequent release of the central regulator Sld3 and other factors associated with single-stranded DNA, leading to a deeper understanding of the initialization mechanism of DNA replication.

Revised version: In conclusion, our structural and biochemical studies of Sld3CBD–Cdc45 revealed a detailed process of CMG formation and the subsequent ssDNA-mediated release of the central regulator Sld3 with other factors, leading to a deeper understanding of the initiation mechanism of DNA replication.

Materials and Methods

Preparation of proteins

The C-terminal His-tagged Sld3CBD of S. cerevisiae was expressed in Escherichia coli and prepared as previously described [23]. For co-expression of S. cerevisiae Sld3CBD (S148–K430) and Cdc45 (M1-L650) (Sld3CBD–Cdc45) in E. coli, the Sld3CBD DNA fragment containing the His6-tag (LEHHHHHH) at the C-terminus and the Cdc45 DNA fragment were cloned in the co-overexpression vector pETDuet-1 between the NcoI and SalI restriction sites and the NdeI and XhoI restriction sites, respectively. The primer sequences for Sld3CBD–Cdc45 are listed in Supplementary Table 1.

To confirm the interactions obtained from the Sld3CBD–Cdc45 structure, we constructed four types of single or multi-mutants for Sld3CBD (Sld3-3S: I352S/I355S/L356S, Sld3-3E: 352E/I355E/L356E, Sld3-2R: D344R/D348R, and Sld3-Y: I352Y), and three types of single or multi-mutants for Cdc45 (Cdc45-IIIS: L522S/L527S/V529S, Cdc45-IIIE: L522E/L527E/V529E, Cdc45-IIS: L637S/L641S, Cdc45-IIE: L637E/L641E, and Cdc45-A: R523A) using the Quick Change site-directed mutagenesis method and the inverse PCR method with pETDuet-1– Sld3CBD–Cdc45 as template DNA. The primer sequences for the mutant strains are listed in Supplementary Table 1.

Earlier version: As S. cerevisiae Sld7–Sld3–Cdc45 could not be co-overexpressed in E. coli, we attempted to clone it from several other fungal sources. The complex of Sld7 (M1-T268), Sld3ΔC (M1-K430, truncated C-terminal domain), and Cdc45 (M1-I666) from the budding yeast Kluyveromyces marxianus, which belongs to the same family as S. cerevisiae (Sld7–Sld3ΔC–Cdc45), was cloned into pETDuet-1. We also added a His6-tag (LEHHHHHH) to the C-terminus of Sld3ΔC. The primer sequences for Sld7–Sld3ΔC and Sld7–Sld3ΔC–Cdc45 are listed in Supplementary Table 1. As Cdc45 mutants can lose the ability to bind to Sld3, we overexpress Sld7–Sld3ΔC by using multi-mutants of Cdc4 (Cdc45-IIS) in Sld7–Sld3ΔC–Cdc45. The Ser-substitution residues (L654S/L658S) of Cdc45-IIS from K. marxianus were selected based on sequence conservation. We constructed mutant Sld7–Sld3ΔC–Cdc45-IIS using the Quick Change site-directed mutagenesis method and the inverse PCR method with pETDuet-1–Sld7–Sld3ΔC–Cdc45 as template DNA. The primer sequences for the mutant strains are listed in Supplementary Table 1.

Revised version: As S. cerevisiae Sld7–Sld3–Cdc45 could not be co-overexpressed in E. coli, we attempted to clone it from several other fungal sources. The complex of Sld7 (M1-T268), Sld3ΔC (M1-K430, truncated C-terminal domain), and Cdc45 (M1-I666) from the budding yeast Kluyveromyces marxianus, which belongs to the same family as S. cerevisiae (Sld7–Sld3ΔC–Cdc45). Sld7, Sld3 and Cdc45 have sequence conservation with similar structures (RMSD = 0.356, 1.392, and 0.891 for Ca atoms of Sld7CTD, Sld7NTD-Sld3NTD, and Sld3CBD-Cdc45) predicted by the Alphafold3 [42]. Sld7–Sld3ΔC–Cdc45 was cloned into pETDuet-1 with a His6-tag (LEHHHHHH) at the C-terminus of Sld3ΔC. The primer sequences for Sld7–Sld3ΔC and Sld7–Sld3ΔC–Cdc45 are listed in Supplementary Table 1. As Cdc45 mutants can lose the ability to bind to Sld3, we overexpress Sld7–Sld3ΔC by using multi-mutants of Cdc45 (Cdc45-IIS) in Sld7–Sld3ΔC–Cdc45. The Ser-substitution residues (L654S/L658S) of Cdc45-IIS from K. marxianus were selected based on sequence conservation. We constructed mutant Sld7–Sld3ΔC–Cdc45-IIS using the Quick Change site-directed mutagenesis method and the inverse PCR method with pETDuet-1–Sld7– Sld3ΔC–Cdc45 as the template DNA. The primer sequences for the mutant strains are listed in Supplementary Table 1.

To overexpress Sld3CBD, Sld3CBD–Cdc45, Sld3CBD–Cdc45 mutants, Sld7–Sld3ΔC–Cdc45 and Sld7–Sld3ΔC, the vector of pET26Sld3CBD, pETDuet-1-Sld3CBD–Cdc45, pETDuet-1-Sld3CBD–Cdc45 mutants, pETDuet-1-Sld7–Sld3ΔC–Cdc45 or pETDuet-1-Sld7–Sld3ΔC was transformed into E. coli strain BL21 (DE3) through electroporation, followed by preculturing in 5 ml of Luria–Bertani medium containing 100 μg/mL ampicillin at 310 K overnight. The culture was then transferred to 3 L Luria–Bertani medium and grown until the OD600 reached 0.6. After cooling the culture for 30 min on ice, overexpression of each sample was induced by the addition of isopropyl-β-D-1-thiogalactopyranoside to a final concentration of 0.5 mM, and then cells were grown for an additional 12 h at 293 K. Cells were harvested by centrifugation at 4,000 g for 20 min at 283 K and then resuspended in a buffer containing 20 mM Tris-HCl pH 7.5, 300 mM NaCl, 10% glycerol, 0.1 mg/ml DNase, and 1× protease-inhibitor cocktail (cOmplete EDTA-free; Roche, Basel, Switzerland).

Earlier version: The harvested cells were crushed through sonication and centrifuged at 40,000 g for 30 min at 283 K. After filtration through a 0.45-µm filter (Sigma-Aldrich/Merck Millipore, Burlington, MA, USA), the supernatant was then loaded onto the HisTrap HP column equilibrated with buffer A [20 mM Tris-HCl, pH 7.5, 300 mM NaCl, 10% glycerol], and the column was washed with buffer A. The His-tagged target protein was eluted with imidazole at 20% (100 mM), 30% (150 mM), and 50% (250 mM), followed by a linear gradient of 250–500 mM imidazole in buffer B [20 mM Tris-HCl, pH 7.5, 300 mM NaCl, 10% glycerol, 500 mM imidazole]. We analyzed the fractions using sodium dodecyl sulfate-polyacrylamide gel electrophoresis (SDS-PAGE). Except for mutants of Sld3CBD–Cdc45, the pooled fractions were then purified through size-exclusion chromatography using a buffer A-equilibrated Superdex 200 16/60 column (GE HealthCare, Chicago, IL, USA). Protein purity was confirmed through SDS-PAGE. The purified samples were concentrated to 10 mg/ml and stored at 193 K.

Revised version: The harvested cells were crushed through sonication and centrifuged at 40,000 g for 30 min at 283 K. After filtration through a 0.45-µm filter (Sigma-Aldrich/Merck Millipore, Burlington, MA, USA), the supernatant was then loaded onto the HisTrap HP column equilibrated with buffer A [20 mM Tris-HCl, pH 7.5, 300 mM NaCl, 10% glycerol] and the column was washed with buffer A. The His-tagged target protein was eluted with imidazole at 20% (100 mM), 30% (150 mM), and 50% (250 mM), followed by a linear gradient of 250–500 mM imidazole in buffer B [20 mM Tris-HCl, pH 7.5, 300 mM NaCl, 10% glycerol, 500 mM imidazole]. We analyzed the fractions using sodium dodecyl sulfate-polyacrylamide gel electrophoresis (SDS-PAGE). Except for mutants of Sld3CBD–Cdc45, the pooled fractions were then purified through size-exclusion chromatography with a buffer A-equilibrated Superdex 200 16/60 column (GE HealthCare, Chicago, IL, USA). Protein purity was confirmed through SDS-PAGE. The purified samples were concentrated to 10 mg/ml and stored at 193 K.

Crystallization and diffraction data collection

The crystallization screening of Sld3CBD–Cdc45 was performed at 293 K using the sitting-drop vapor diffusion method. The drop was a mixture of 1.0 µl of a 10 mg/ml protein solution with the equivalent volume of reservoir buffer from the commercial crystallization kits (JCSG core I-IV, classics, classics II, PEGs, PEGsII, and MPD suite from Qiagen, Venlo, the Netherlands). The initial needle- and plate-shaped crystals appeared under eight and one conditions, respectively. After optimizing the conditions by altering the type and concentration of precipitant, the salt reagent, and the pH of the buffer, the best crystals were obtained in a solution of 0.2M sodium acetate, 0.1M Bis-Tris propane (pH 6.5: pH 8.5 = 3:7), 20%(w/v) PEG3500, and a drop containing 2.0 µl of a 10 mg/ml protein solution mixed with the equivalent volume of reservoir buffer.

Earlier version: Diffraction data were collected using a beamline BL-17A at the Photon Factory, Tsukuba, Japan. Before data collection, the crystals were cryoprotected by soaking them in a reservoir buffer supplemented with 5% (v/v) glycerol and then flash-cooled under a nitrogen gas stream at 100 K. XDS/XSCALE was used to index, integrate, scale, and merge the dataset [42]. Supplementary Table 2 summarizes the data collection and data processing statistics.

Revised version: Diffraction data were collected using a beamline BL-17A at the Photon Factory, Tsukuba, Japan. Before data collection, the crystals were cryoprotected by soaking them in a reservoir buffer supplemented with 5% (v/v) glycerol, and then flash-cooled under a nitrogen gas stream at 100 K. The XDS/XSCALE program was used to index, integrate, scale, and merge the dataset [43]. Supplementary Table 2 summarizes the statistics of data collection and processing.

Structure determination and refinement

Earlier version: The structure of Sld3CBD–Cdc45 was determined using the molecular replacement method with the AutoMR wizard in Phenix [43, 44] and human Cdc45 (PDBID: 5DGO) [31] and Sld3CBD (PDBID:3WI3) [23] structures were used as search models. Several rounds of refinement were performed using the Phenix_refine program in Phenix, interleaved with manual building and fitting according to the electron density maps of 2Fo-Fc and Fo-Fc using the Coot program [44, 45]. Supplementary Table 2 shows the final refinement statistics and geometry defined by MolProbity [46]. All structural diagrams were generated using PyMol [47].

Revised version: The structure of Sld3CBD–Cdc45 was determined using the molecular replacement method with the AutoMR wizard in Phenix [44, 45]. The structures of human Cdc45 (PDBID: 5DGO) [31] and Sld3CBD (PDBID: 3WI3) [23] were used as search models. Several rounds of refinement were performed using the Phenix_refine program in Phenix, interleaved with manual building and fitting according to the electron density maps of 2Fo-Fc and Fo-Fc using the Coot program [45, 46]. Supplementary Table 2 presents the final refinement statistics and geometry defined by MolProbity [47]. All structural diagrams were generated using PyMol [48].

Mutant analysis of Sld3 and Cdc45

Earlier version: To analyze the binding sites of Sld3CBD-Cdc45, in conjunction with Cdc45 binding sites to MCM and GINS, we performed co-express pull-down binding assay. We constructed four variants of Sld3CBD and five variants of Cdc45 according to the binding information from the obtained Sld3CBD–Cdc45 structure. We co-express all mutations of Sld3CBD-Cdc45 (Sld3CBD-C-histag) and load them onto the HisTrap HP column under the same conditions as in [Preparation of proteins]. After extracting the samples through Ni-affinity chromatography, we concentrated each eluted sample to 0.005 mg/mL (by nanodrop A280) by centrifugation and checked the binding status of the mutants using SDS-PAGE. Both Sld3CBD and Cdc45 should be observed in the elution group if they formed a complex. The overexpress level of the Cdc45 was checked by - IPTG and +IPTG.

Revised version: To analyze the binding sites of Sld3CBD-Cdc45, in conjunction with Cdc45 binding sites to MCM and GINS, we performed a co-express pull-down binding assay. We constructed four variants of Sld3CBD and five variants of Cdc45 according to the binding information from our Sld3CBD–Cdc45 structure. We co-express all mutations of Sld3CBD-Cdc45 (Sld3CBD-C-histag) and load them onto the HisTrap HP column under the same conditions as in [Preparation of proteins]. After extracting the samples through Ni-affinity chromatography, we concentrated each eluted sample to 0.005 mg/mL (by nanodrop A280) and checked the binding status of the mutants using SDS-PAGE. Both Sld3CBD and Cdc45 should be observed in the elution group if they form a complex. The overexpressed level of the Cdc45 was checked by -IPTG and +IPTG.

Earlier version: To check whether mutations affected the structure, we performed circular dichroism (CD) spectrometry measurement of Sld3mutants–Cdc45. CD spectra were collected using a J-805 spectropolarimeter (JASCO, Tokyo, Japan) in a quartz cell with a path length of 1 mm in an atmosphere of N2 at 298K. For CD measurements, the samples were dialyzed in a buffer [20 mM Tris-HCl, pH 7.5, 50 mM NaCl] and adjusted to a concentration of 0.5 mg/mL through absorption. CD spectra for the wavelength range of 190–300 nm were obtained by averaging the results of four scans. The results are given in molar ellipticity per residues [θ] mrw (× 10−3) vs the wavelength/nm [48]. The secondary structures of each sample were estimated using the K2D3 method [49]. For wild-type Sld3CBD, we calculated secondary structures from the obtained Sld3CBD–Cdc45 structure.

Revised version: We performed circular dichroism (CD) spectrometry measurement of Sld3 mutants–Cdc45 to check whether mutations affected the structure. CD spectra were collected using a J-805 spectropolarimeter (JASCO, Tokyo, Japan) in a quartz cell with a path length of 1 mm in an atmosphere of N2 at 298 K. For CD measurements, the samples were dialyzed in a buffer [20 mM Tris-HCl, pH 7.5, 50 mM NaCl] and adjusted to a concentration of 0.5 mg/mL through absorption. CD spectra for the wavelength range of 190–300 nm were obtained by averaging the results of four scans. The results are given in molar ellipticity per residue [θ] mrw (× 10−3) vs wavelength/nm [49]. The secondary structures of each sample were estimated using the K2D3 method [50]. For wild-type Sld3CBD, we calculated secondary structures from the obtained Sld3CBD–Cdc45 structure.

Growth of mutant cells

The isolation and plasmid shuffling of the temperature-sensitive yeast strain YYK13 (for Sld3 mutant analysis) and Cdc45-35 (for Cdc45 mutant analysis) have been described in a previous study [12]. To investigate Sld3 mutants, strain YYK13 was transformed with the YCplac22 plasmid containing multiple variants of SLD3 mutant genes. The transformants were streaked on SD plates lacking Trp and Leu (SD-Trp, Leu) and SD-Trp, Leu containing 5-fluoroorotic acid (5-FOA-Trp, Leu), and then incubated at 298 K for 3 days. To analyze the Cdc45 mutant, strain Cdc45-35 containing the CDC45 mutant gene was re-transformed with the YEplac195 plasmid containing SLD3 and the YEplac122 plasmid containing SLD7. The transformants were streaked on yeast extract–peptone–dextrose plates and incubated at 298 or 305 K for 3 days.

Modelling of complexes

We constructed the models of Sld3CBD–Cdc45–MCM–dsDNA (Sld3CBD–Cdc45–MCM dimer complexed with dsDNA) and SCMG–dsDNA (SCMG dimer complexed with dsDNA) using a two-step process. First, the structure of the MCM–NTDs in the CMG monomer (PDBID: 3JC6) [28] was superimposed on that of MCM DH complexed with dsDNA (5BK4) [35] to obtain a dimer of Cdc45–MCM–GINS complexed with dsDNA (Cdc45–MCM–GINS–dsDNA). Next, the models of Sld3CBD–Cdc45–MCM–dsDNA and SCMG–dsDNA were obtained by superimposing Cdc45NTD from Sld3CBD–Cdc45 onto Cdc45–MCM–GINS–dsDNA.

Dynamic light scattering

Earlier version: We estimated the particle size of the Sld7–Sld3ΔC–Cdc45 complex using dynamic light scattering in terms of peak size [50]. To avoid the effect of concentration, we measured four samples of the Sld7–Sld3ΔC–Cdc45 complex. Each sample was overexpressed independently and purified through size-exclusion chromatography. All samples were concentrated to 10 mg/ml, and stored at 193 K. Before measurement, protein samples were dialyzed in a buffer containing 20 mM Tris-HCl pH 7.5 and 300 mM NaCl and then filtered through a 0.22-μm filter. The pure buffer was measured as a background control and 10 μM lysozyme was measured as a standard control. Dynamic light scattering data were collected and analyzed using a Malvern Zetasizer Nano-ZS instrument (Malvern Panalytical, Malvern, UK) for 0.2–0.5 mg/ml protein samples 500 μl in a microquartz cuvette (Malvern-ZEN0112). An automatic duration model was used to collect data. The Zetasizer 6.20 was utilized for data analysis from three measurements to estimate the particle size of each sample.

Revised version: We estimated the particle size of the Sld7–Sld3ΔC–Cdc45 complex using dynamic light scattering in terms of peak size [51]. To avoid the effect of concentration, we measured four samples of the Sld7–Sld3ΔC–Cdc45 complex. Each sample was overexpressed independently and purified through size-exclusion chromatography. All samples were concentrated to 10 mg/ml, and stored at 193 K. Before measurement, protein samples were dialyzed in a buffer containing 20 mM Tris-HCl pH 7.5 and 300 mM NaCl and then filtered through a 0.22-μm filter. The pure buffer was measured as a background control and 10 μM lysozyme was measured as a standard control. Dynamic light scattering data were collected and analyzed using a Malvern Zetasizer Nano-ZS instrument (Malvern Panalytical, Malvern, UK) for 0.2–0.5 mg/ml, 500 μl protein samples in a microquartz cuvette (Malvern-ZEN0112). An automatic duration model was used to collect data. The Zetasizer 6.20 was utilized for data analysis from three measurements to estimate the particle size of each sample.

Electrophoretic mobility shift assay for ssDNA binding

Earlier version: Sld3 binds to single strands of DNA (ssARS1-2 and ssARS1-5) [26]. To determine the specificity of the ssDNA binding affinity of Sld3, we conducted an electrophoretic mobility shift assay (EMSA) using non-denaturing PAGE (native-PAGE) [51]. Given that Cdc45 binds ssDNA with a nonspecific sequence at lengths greater than 60 bases [36], we designed three fragments of ssDNA 40 bases in length (first half 1–40 bp, second half 41–80 bp, and middle half 21–60 bp) for each ssDNA 80-bp segment. All ssDNA fragments of S. cerevisiae were synthesized by Sigma-Aldrich. The 40-bp dsDNA fragments (dsARS1-34_1:ssARS1-3_1/ssARS1-4_2) were converted by annealing them using the PCR protocol and then checked through polyacrylamide gel electrophoresis. The loaded samples were incubated overnight at 277 K in TMK buffer (20mM Tris-HCl pH8, 100mM MgCl2, 200mM KCl) containing synthesized ssDNA and varying amounts of proteins at ssDNA: protein molecular ratios of 1:0, 1:1, 1:2, 1:4, and 0:1 with 10 pM as 1 unit. After incubation, the mixtures were loaded onto a polyacrylamide gel (5% (w/v) acrylamide (39:1), 10% 10 × running buffer (0.25 M Tris, 1.92 M Glycine), 0.1% Ammonium peroxodisulphate, 0.06% (v/v) TEMED) without denaturing (native-PAGE). EMSA was performed at 10 mA/200 V per gel for 40 min at 277 K in 1 × running buffer. After electrophoresis, the reaction products were visualized using SYBR safe (SYBR safe:1 × running buffer = 0.0001:1) to stain the ssDNA and Coomassie brilliant blue (G-250) to stain the proteins. We repeated the EMSA experiments three or more times to confirm the readability. Considering the functional similarity of ARS1-core, the EMSA of Sld7–Sld3ΔC and Sld7–Sld3ΔC–Cdc45 of K. marxianus used ssDNA fragments of S. cerevisiae.

Revised version: Sld3 binds to single strands of DNA (ssARS1-2 and ssARS1-5) [26]. To determine the specificity of the ssDNA binding affinity of Sld3, we conducted an electrophoretic mobility shift assay (EMSA) using non-denaturing PAGE (native-PAGE) [52]. Given that Cdc45 binds ssDNA with a nonspecific sequence at lengths greater than 60 bases [36], we designed three fragments of ssDNA 40 bases in length (first half 1–40 bp, second half 41–80 bp, and middle half 21–60 bp) for each ssDNA 80-bp segment. All ssDNA fragments of S. cerevisiae were synthesized by Sigma-Aldrich. The 40-bp dsDNA fragments (dsARS1-34_1:ssARS1-3_1/ssARS1-4_2) were converted by annealing them using the PCR protocol and then checked through polyacrylamide gel electrophoresis. The loaded samples were incubated overnight at 277 K in TMK buffer (20mM Tris-HCl pH8, 100mM MgCl2, 200mM KCl) containing synthesized ssDNA and varying amounts of proteins at ssDNA: protein molecular ratios of 1:0, 1:0.5, 1:1, 1:2, and 0:1 with 20 pM as 1 unit. After incubation, the mixtures were loaded onto a polyacrylamide gel (5% (w/v) acrylamide (39:1), 10% 10 × running buffer (0.25 M Tris, 1.92 M Glycine), 0.1% Ammonium peroxodisulphate, 0.06% (v/v) TEMED) without denaturing (native-PAGE). EMSA was performed at 10 mA/200 V per gel for 40 min at 277 K in 1 × running buffer. After electrophoresis, the reaction products were visualized using Fast Blast DNA stain (Bio-Rad Laboratories, Hercules, CA, USA) (100× Fast Blast DNA stain diluted by 1 × running buffer) or SYBR safe (Sigma-Aldrich, Burlington, MA, USA) (SYBR safe:1 × running buffer = 0.0001:1) to stain the ssDNA and Coomassie brilliant blue (G-250) to stain the proteins. We repeated the EMSA experiments three or more times to confirm the readability. All EMSA results were converted into 8-bit images after brightness and contrast normalized, then the background was removed, and the integrated grayscale was calculated using imageJ [53]. The results of band-integrated grayscale calculation were performed with t-test and plotted using GraphPAD Prism (Graphpad software, San Diego, CA, USA). Considering the functional similarity of ARS1-core, the EMSA of Sld7–Sld3ΔC and Sld7–Sld3ΔC–Cdc45 of K. marxianus used ssDNA fragments of S. cerevisiae.